(assume a0 (! (not (! (= (! (bvlshr s (! (bvor (! (bvand t #b0000) :named @p_1) s) :named @p_2)) :named @p_3) #b0000) :named @p_4)) :named @p_5))
(step t0 (cl (not (! (= (! (or (! (not (! (not (! (= (! (ite (! (= (! (ubv_to_int s) :named @p_6) 3) :named @p_16) (! (div @p_6 8) :named @p_15) (! (ite (! (= @p_6 2) :named @p_13) (! (div @p_6 4) :named @p_12) (! (ite (! (= @p_6 1) :named @p_10) (! (div @p_6 2) :named @p_9) (! (ite (! (= @p_6 0) :named @p_7) @p_6 0) :named @p_8)) :named @p_11)) :named @p_14)) :named @p_17) 0) :named @p_18)) :named @p_19)) :named @p_122) (! (not (! (= @p_17 @p_14) :named @p_79)) :named @p_312) (! (not (! (= @p_14 @p_11) :named @p_252)) :named @p_486) (! (not (! (= @p_11 @p_9) :named @p_433)) :named @p_533) (! (not (! (>= (! (+ (! (ubv_to_int s) :named @p_6) (! (* -2 @p_9) :named @p_55)) :named @p_56) 0) :named @p_57)) :named @p_650) (! (not (! (>= @p_17 0) :named @p_532)) :named @p_535) (! (not (! (not (! (>= @p_6 2) :named @p_514)) :named @p_515)) :named @p_517)) :named @p_652) (! (or @p_18 @p_312 @p_486 @p_533 @p_650 @p_535 @p_514) :named @p_651)) :named @p_712)) (not @p_652) @p_651) :rule equiv_pos2)
(step t1 (cl (! (= @p_122 @p_18) :named @p_213)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_213 1 6))
(step t2 (cl (= @p_312 @p_312)) :rule refl)
(step t3 (cl (= @p_486 @p_486)) :rule refl)
(step t4 (cl (= @p_533 @p_533)) :rule refl)
(step t5 (cl (= @p_650 @p_650)) :rule refl)
(step t6 (cl (= @p_535 @p_535)) :rule refl)
(step t7 (cl (! (= @p_517 @p_514) :named @p_713)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_713 1 6))
(step t8 (cl @p_712) :rule cong :premises (t1 t2 t3 t4 t5 t6 t7))
(step t9 (cl (! (=> (! (and @p_19 @p_79 @p_252 @p_433 @p_57 @p_532 @p_515) :named @p_654) (! (and @p_433 @p_252 @p_515 @p_19 @p_532 @p_79 @p_57) :named @p_656)) :named @p_711) @p_654) :rule implies_neg1)
(anchor :step t10)
(assume t10.a0 @p_19)
(assume t10.a1 @p_79)
(assume t10.a2 @p_252)
(assume t10.a3 @p_433)
(assume t10.a4 @p_57)
(assume t10.a5 @p_532)
(assume t10.a6 @p_515)
(step t10.t0 (cl @p_656 @p_533 @p_486 @p_517 @p_122 @p_535 @p_312 @p_650) :rule and_neg)
(step t10.t1 (cl (not (! (= @p_5 @p_19) :named @p_20)) (not @p_5) @p_19) :rule equiv_pos2)
(step t10.t2 (cl (! (= @p_4 (! (= #b0000 @p_3) :named @p_48)) :named @p_54)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_54 4 6))
(step t10.t3 (cl (= #b0000 #b0000)) :rule refl)
(step t10.t4 (cl (= s s)) :rule refl)
(step t10.t5 (cl (! (= @p_2 (! (bvor s @p_1) :named @p_49)) :named @p_53)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_53 4 6))
(step t10.t6 (cl (! (= @p_1 #b0000) :named @p_52)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_52 4 6))
(step t10.t7 (cl (= @p_49 (! (bvor s #b0000) :named @p_50))) :rule cong :premises (t10.t4 t10.t6))
(step t10.t8 (cl (! (= @p_50 s) :named @p_51)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_51 4 7))
(step t10.t9 (cl (= @p_49 s)) :rule trans :premises (t10.t7 t10.t8))
(step t10.t10 (cl (= @p_2 s)) :rule trans :premises (t10.t5 t10.t9))
(step t10.t11 (cl (= @p_3 (! (bvlshr s s) :named @p_45))) :rule cong :premises (t10.t4 t10.t10))
(step t10.t12 (cl (= @p_48 (! (= #b0000 @p_45) :named @p_46))) :rule cong :premises (t10.t3 t10.t11))
(step t10.t13 (cl (= @p_4 @p_46)) :rule trans :premises (t10.t2 t10.t12))
(step t10.t14 (cl (= @p_5 (! (not @p_46) :named @p_47))) :rule cong :premises (t10.t13))
(step t10.t15 (cl (= @p_47 (! (not (! (= 0 (! (ite @p_16 @p_15 (! (ite @p_13 @p_12 (! (ite @p_10 @p_9 (! (ite @p_7 (! (div @p_6 1) :named @p_35) 0) :named @p_36)) :named @p_37)) :named @p_38)) :named @p_39)) :named @p_40)) :named @p_41))) :rule bv_intblast)
(step t10.t16 (cl (= 0 0)) :rule refl)
(step t10.t17 (cl (! (= @p_16 @p_16) :named @p_34)) :rule refl)
(step t10.t18 (cl (! (= @p_15 @p_15) :named @p_33)) :rule refl)
(step t10.t19 (cl (! (= @p_13 @p_13) :named @p_32)) :rule refl)
(step t10.t20 (cl (! (= @p_12 @p_12) :named @p_31)) :rule refl)
(step t10.t21 (cl (! (= @p_10 @p_10) :named @p_30)) :rule refl)
(step t10.t22 (cl (! (= @p_9 @p_9) :named @p_29)) :rule refl)
(step t10.t23 (cl (! (= @p_7 @p_7) :named @p_28)) :rule refl)
(step t10.t24 (cl (! (= @p_35 @p_6) :named @p_44)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_44 3 7))
(step t10.t25 (cl (= @p_36 @p_8)) :rule cong :premises (t10.t23 t10.t24 t10.t16))
(step t10.t26 (cl (= @p_37 @p_11)) :rule cong :premises (t10.t21 t10.t22 t10.t25))
(step t10.t27 (cl (= @p_38 @p_14)) :rule cong :premises (t10.t19 t10.t20 t10.t26))
(step t10.t28 (cl (= @p_39 @p_17)) :rule cong :premises (t10.t17 t10.t18 t10.t27))
(step t10.t29 (cl (= @p_40 (! (= 0 @p_17) :named @p_42))) :rule cong :premises (t10.t16 t10.t28))
(step t10.t30 (cl (! (= @p_42 @p_18) :named @p_43)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_43 3 7))
(step t10.t31 (cl (= @p_40 @p_18)) :rule trans :premises (t10.t29 t10.t30))
(step t10.t32 (cl (= @p_41 @p_19)) :rule cong :premises (t10.t31))
(step t10.t33 (cl (= (! (= (! (= @p_17 @p_17) :named @p_21) true) :named @p_22) @p_21)) :rule equiv_simplify)
(step t10.t34 (cl (not @p_22) @p_21) :rule equiv1 :premises (t10.t33))
(step t10.t35 (cl (! (= @p_21 @p_21) :named @p_23)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_23 3 7))
(step t10.t36 (cl @p_21) :rule refl)
(step t10.t37 (cl (! (= @p_6 @p_6) :named @p_27)) :rule refl)
(step t10.t38 (cl (= 3 3)) :rule refl)
(step t10.t39 (cl @p_34) :rule cong :premises (t10.t37 t10.t38))
(step t10.t40 (cl (= 8 8)) :rule refl)
(step t10.t41 (cl @p_33) :rule cong :premises (t10.t37 t10.t40))
(step t10.t42 (cl (= 2 2)) :rule refl)
(step t10.t43 (cl @p_32) :rule cong :premises (t10.t37 t10.t42))
(step t10.t44 (cl (= 4 4)) :rule refl)
(step t10.t45 (cl @p_31) :rule cong :premises (t10.t37 t10.t44))
(step t10.t46 (cl (= 1 1)) :rule refl)
(step t10.t47 (cl @p_30) :rule cong :premises (t10.t37 t10.t46))
(step t10.t48 (cl @p_29) :rule cong :premises (t10.t37 t10.t42))
(step t10.t49 (cl @p_28) :rule cong :premises (t10.t37 t10.t16))
(step t10.t50 (cl (! (= @p_8 @p_8) :named @p_26)) :rule cong :premises (t10.t49 t10.t37 t10.t16))
(step t10.t51 (cl (! (= @p_11 @p_11) :named @p_25)) :rule cong :premises (t10.t47 t10.t48 t10.t50))
(step t10.t52 (cl (! (= @p_14 @p_14) :named @p_24)) :rule cong :premises (t10.t43 t10.t45 t10.t51))
(step t10.t53 (cl @p_21) :rule cong :premises (t10.t39 t10.t41 t10.t52))
(step t10.t54 (cl @p_21) :rule trans :premises (t10.t36 t10.t53))
(step t10.t55 (cl @p_23) :rule cong :premises (t10.t54 t10.t53))
(step t10.t56 (cl @p_22) :rule hole :args ("TRUST_THEORY_REWRITE" @p_22 3 6))
(step t10.t57 (cl @p_22) :rule trans :premises (t10.t35 t10.t55 t10.t56))
(step t10.t58 (cl @p_21) :rule resolution :premises (t10.t34 t10.t57))
(step t10.t59 (cl (= @p_18 @p_18)) :rule cong :premises (t10.t58 t10.t16))
(step t10.t60 (cl (= @p_19 @p_19)) :rule cong :premises (t10.t59))
(step t10.t61 (cl @p_20) :rule trans :premises (t10.t14 t10.t15 t10.t32 t10.t60))
(step t10.t62 (cl @p_19) :rule resolution :premises (t10.t1 t10.t61 a0))
(step t10.t63 (cl (not (! (= (! (and (! (<= (! (* 2 @p_9) :named @p_64) @p_6) :named @p_65) (! (< @p_6 (! (* 2 (! (+ @p_9 1) :named @p_61)) :named @p_62)) :named @p_63)) :named @p_66) (! (and @p_57 (! (not (! (>= @p_56 2) :named @p_58)) :named @p_59)) :named @p_60)) :named @p_67)) (not @p_66) @p_60) :rule equiv_pos2)
(step t10.t64 (cl (! (= @p_65 @p_57) :named @p_78)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_78 3 7))
(step t10.t65 (cl (! (= @p_63 (! (not (! (>= @p_6 @p_62) :named @p_68)) :named @p_69)) :named @p_77)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_77 3 6))
(step t10.t66 (cl @p_27) :rule refl)
(step t10.t67 (cl (! (= @p_61 (! (+ 1 @p_9) :named @p_73)) :named @p_76)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_76 3 7))
(step t10.t68 (cl (= @p_62 (! (* 2 @p_73) :named @p_74))) :rule cong :premises (t10.t42 t10.t67))
(step t10.t69 (cl (! (= @p_74 (! (+ 2 @p_64) :named @p_70)) :named @p_75)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_75 3 7))
(step t10.t70 (cl (= @p_62 @p_70)) :rule trans :premises (t10.t68 t10.t69))
(step t10.t71 (cl (= @p_68 (! (>= @p_6 @p_70) :named @p_71))) :rule cong :premises (t10.t66 t10.t70))
(step t10.t72 (cl (! (= @p_71 @p_58) :named @p_72)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_72 3 7))
(step t10.t73 (cl (= @p_68 @p_58)) :rule trans :premises (t10.t71 t10.t72))
(step t10.t74 (cl (= @p_69 @p_59)) :rule cong :premises (t10.t73))
(step t10.t75 (cl (= @p_63 @p_59)) :rule trans :premises (t10.t65 t10.t74))
(step t10.t76 (cl @p_67) :rule cong :premises (t10.t64 t10.t75))
(step t10.t77 (cl (and @p_29 @p_66)) :rule hole :args ("ARITH_REDUCTION" @p_9))
(step t10.t78 (cl @p_66) :rule and :premises (t10.t77) :args (1))
(step t10.t79 (cl @p_60) :rule resolution :premises (t10.t63 t10.t76 t10.t78))
(step t10.t80 (cl @p_57) :rule and :premises (t10.t79) :args (0))
(step t10.t81 (cl @p_656) :rule resolution :premises (t10.t0 t10.a3 t10.a2 t10.a6 t10.t62 t10.a5 t10.a1 t10.t80))
(step t10 (cl @p_122 @p_312 @p_486 @p_533 @p_650 @p_535 @p_517 @p_656) :rule subproof :discharge (t10.a0 t10.a1 t10.a2 t10.a3 t10.a4 t10.a5 t10.a6))
(step t11 (cl (! (not @p_654) :named @p_655) @p_19) :rule and_pos :args (0))
(step t12 (cl @p_655 @p_79) :rule and_pos :args (1))
(step t13 (cl @p_655 @p_252) :rule and_pos :args (2))
(step t14 (cl @p_655 @p_433) :rule and_pos :args (3))
(step t15 (cl @p_655 @p_57) :rule and_pos :args (4))
(step t16 (cl @p_655 @p_532) :rule and_pos :args (5))
(step t17 (cl @p_655 @p_515) :rule and_pos :args (6))
(step t18 (cl @p_656 @p_655 @p_655 @p_655 @p_655 @p_655 @p_655 @p_655) :rule resolution :premises (t10 t11 t12 t13 t14 t15 t16 t17))
(step t19 (cl @p_655 @p_655 @p_655 @p_655 @p_655 @p_655 @p_655 @p_656) :rule reordering :premises (t18))
(step t20 (cl @p_655 @p_656) :rule contraction :premises (t19))
(step t21 (cl @p_711 @p_656) :rule resolution :premises (t9 t20))
(step t22 (cl @p_711 (! (not @p_656) :named @p_657)) :rule implies_neg2)
(step t23 (cl @p_711 @p_711) :rule resolution :premises (t21 t22))
(step t24 (cl @p_711) :rule contraction :premises (t23))
(step t25 (cl @p_655 @p_656) :rule implies :premises (t24))
(step t26 (cl (! (=> @p_656 false) :named @p_658) @p_656) :rule implies_neg1)
(anchor :step t27)
(assume t27.a0 @p_433)
(assume t27.a1 @p_252)
(assume t27.a2 @p_515)
(assume t27.a3 @p_19)
(assume t27.a4 @p_532)
(assume t27.a5 @p_79)
(assume t27.a6 @p_57)
(step t27.t0 (cl (not @p_67) (not @p_66) @p_60) :rule equiv_pos2)
(step t27.t1 (cl @p_78) :rule hole :args ("TRUST_THEORY_REWRITE" @p_78 3 7))
(step t27.t2 (cl @p_77) :rule hole :args ("TRUST_THEORY_REWRITE" @p_77 3 6))
(step t27.t3 (cl @p_27) :rule refl)
(step t27.t4 (cl (= 2 2)) :rule refl)
(step t27.t5 (cl @p_76) :rule hole :args ("TRUST_THEORY_REWRITE" @p_76 3 7))
(step t27.t6 (cl (= @p_62 @p_74)) :rule cong :premises (t27.t4 t27.t5))
(step t27.t7 (cl @p_75) :rule hole :args ("TRUST_THEORY_REWRITE" @p_75 3 7))
(step t27.t8 (cl (= @p_62 @p_70)) :rule trans :premises (t27.t6 t27.t7))
(step t27.t9 (cl (= @p_68 @p_71)) :rule cong :premises (t27.t3 t27.t8))
(step t27.t10 (cl @p_72) :rule hole :args ("TRUST_THEORY_REWRITE" @p_72 3 7))
(step t27.t11 (cl (= @p_68 @p_58)) :rule trans :premises (t27.t9 t27.t10))
(step t27.t12 (cl (= @p_69 @p_59)) :rule cong :premises (t27.t11))
(step t27.t13 (cl (= @p_63 @p_59)) :rule trans :premises (t27.t2 t27.t12))
(step t27.t14 (cl @p_67) :rule cong :premises (t27.t1 t27.t13))
(step t27.t15 (cl (and @p_29 @p_66)) :rule hole :args ("ARITH_REDUCTION" @p_9))
(step t27.t16 (cl @p_66) :rule and :premises (t27.t15) :args (1))
(step t27.t17 (cl @p_60) :rule resolution :premises (t27.t0 t27.t14 t27.t16))
(step t27.t18 (cl @p_57) :rule and :premises (t27.t17) :args (0))
(step t27.t19 (cl (not (! (= (! (< @p_56 0) :named @p_659) @p_650) :named @p_710)) (not @p_659) @p_650) :rule equiv_pos2)
(step t27.t20 (cl @p_710) :rule hole :args ("TRUST_THEORY_REWRITE" @p_710 3 6))
(step t27.t21 (cl (not (! (= @p_650 @p_659) :named @p_709)) (! (not @p_650) :named @p_653) @p_659) :rule equiv_pos2)
(step t27.t22 (cl @p_709) :rule symm :premises (t27.t20))
(step t27.t23 (cl (! (=> @p_57 false) :named @p_660) @p_57) :rule implies_neg1)
(anchor :step t27.t24)
(assume t27.t24.a0 @p_57)
(step t27.t24.t0 (cl (not (! (= (! (< (! (+ (! (* -1 @p_56) :named @p_668) (! (* 2 (! (+ @p_17 (! (* -1 @p_14) :named @p_333)) :named @p_334)) :named @p_667) (! (* -2 @p_17) :named @p_666) @p_6 (! (* 2 (! (+ @p_14 (! (* -1 @p_11) :named @p_557)) :named @p_558)) :named @p_665) (! (* 2 (! (+ @p_11 (! (* -1 @p_9) :named @p_547)) :named @p_548)) :named @p_664)) :named @p_669) (! (+ (! (* -1 0) :named @p_136) (! (* 2 0) :named @p_661) (! (* -2 1) :named @p_662) 2 @p_661 @p_661) :named @p_663)) :named @p_670) false) :named @p_691)) (not @p_670) false) :rule equiv_pos2)
(step t27.t24.t1 (cl (! (= @p_670 (! (not (! (>= @p_669 @p_663) :named @p_692)) :named @p_693)) :named @p_708)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_708 3 6))
(step t27.t24.t2 (cl (! (= @p_668 (! (+ (! (* -1 @p_6) :named @p_140) @p_64) :named @p_701)) :named @p_707)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_707 3 7))
(step t27.t24.t3 (cl (! (= @p_667 (! (+ (* 2 @p_17) (* -2 @p_14)) :named @p_700)) :named @p_706)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_706 3 7))
(step t27.t24.t4 (cl (= @p_666 @p_666)) :rule refl)
(step t27.t24.t5 (cl (! (= @p_665 (! (+ (* 2 @p_14) (* -2 @p_11)) :named @p_699)) :named @p_705)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_705 3 7))
(step t27.t24.t6 (cl (! (= @p_664 (! (+ (* 2 @p_11) @p_55) :named @p_698)) :named @p_704)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_704 3 7))
(step t27.t24.t7 (cl (= @p_669 (! (+ @p_701 @p_700 @p_666 @p_6 @p_699 @p_698) :named @p_702))) :rule cong :premises (t27.t24.t2 t27.t24.t3 t27.t24.t4 t27.t3 t27.t24.t5 t27.t24.t6))
(step t27.t24.t8 (cl (! (= @p_702 0) :named @p_703)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_703 3 7))
(step t27.t24.t9 (cl (= @p_669 0)) :rule trans :premises (t27.t24.t7 t27.t24.t8))
(step t27.t24.t10 (cl (! (= @p_136 0) :named @p_199)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_199 3 6))
(step t27.t24.t11 (cl (! (= @p_661 0) :named @p_696)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_696 3 6))
(step t27.t24.t12 (cl (! (= @p_662 -2) :named @p_697)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_697 3 7))
(step t27.t24.t13 (cl (= @p_663 (! (+ 0 0 -2 2 0 0) :named @p_694))) :rule cong :premises (t27.t24.t10 t27.t24.t11 t27.t24.t12 t27.t4 t27.t24.t11 t27.t24.t11))
(step t27.t24.t14 (cl (! (= @p_694 0) :named @p_695)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_695 3 7))
(step t27.t24.t15 (cl (= @p_663 0)) :rule trans :premises (t27.t24.t13 t27.t24.t14))
(step t27.t24.t16 (cl (= @p_692 (! (>= 0 0) :named @p_195))) :rule cong :premises (t27.t24.t9 t27.t24.t15))
(step t27.t24.t17 (cl (! (= @p_195 true) :named @p_196)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_196 3 7))
(step t27.t24.t18 (cl (= @p_692 true)) :rule trans :premises (t27.t24.t16 t27.t24.t17))
(step t27.t24.t19 (cl (= @p_693 (! (not true) :named @p_193))) :rule cong :premises (t27.t24.t18))
(step t27.t24.t20 (cl (! (= @p_193 false) :named @p_194)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_194 1 7))
(step t27.t24.t21 (cl (= @p_693 false)) :rule trans :premises (t27.t24.t19 t27.t24.t20))
(step t27.t24.t22 (cl @p_691) :rule trans :premises (t27.t24.t1 t27.t24.t21))
(step t27.t24.t23 (cl (not (! (<= @p_668 @p_136) :named @p_689)) (not (! (= @p_667 @p_661) :named @p_687)) (not (! (<= @p_666 @p_662) :named @p_679)) (! (not (! (< @p_6 2) :named @p_518)) :named @p_531) (not (! (= @p_665 @p_661) :named @p_676)) (not (! (= @p_664 @p_661) :named @p_671)) @p_670) :rule la_generic :args (1/1 -1/1 1/1 1/1 -1/1 -1/1 1/1))
(step t27.t24.t24 (cl (=> (! (and (! (< -1 0) :named @p_147) @p_57) :named @p_690) @p_689)) :rule la_mult_neg)
(step t27.t24.t25 (cl (not @p_690) @p_689) :rule implies :premises (t27.t24.t24))
(step t27.t24.t26 (cl @p_690 (! (not @p_147) :named @p_150) @p_650) :rule and_neg)
(step t27.t24.t27 (cl (= (! (= @p_147 true) :named @p_149) @p_147)) :rule equiv_simplify)
(step t27.t24.t28 (cl (not @p_149) @p_147) :rule equiv1 :premises (t27.t24.t27))
(step t27.t24.t29 (cl @p_149) :rule rare_rewrite :args ("evaluate"))
(step t27.t24.t30 (cl @p_147) :rule resolution :premises (t27.t24.t28 t27.t24.t29))
(step t27.t24.t31 (cl @p_690) :rule resolution :premises (t27.t24.t26 t27.t24.t30 t27.t18))
(step t27.t24.t32 (cl @p_689) :rule resolution :premises (t27.t24.t25 t27.t24.t31))
(step t27.t24.t33 (cl (=> (! (and (! (> 2 0) :named @p_672) (! (= @p_334 0) :named @p_355)) :named @p_688) @p_687)) :rule la_mult_pos)
(step t27.t24.t34 (cl (not @p_688) @p_687) :rule implies :premises (t27.t24.t33))
(step t27.t24.t35 (cl @p_688 (! (not @p_672) :named @p_675) (! (not @p_355) :named @p_361)) :rule and_neg)
(step t27.t24.t36 (cl (= (! (= @p_672 true) :named @p_674) @p_672)) :rule equiv_simplify)
(step t27.t24.t37 (cl (not @p_674) @p_672) :rule equiv1 :premises (t27.t24.t36))
(step t27.t24.t38 (cl @p_674) :rule rare_rewrite :args ("evaluate"))
(step t27.t24.t39 (cl @p_672) :rule resolution :premises (t27.t24.t37 t27.t24.t38))
(step t27.t24.t40 (cl (! (not (! (= @p_79 @p_355) :named @p_358)) :named @p_360) @p_312 @p_355) :rule equiv_pos2)
(step t27.t24.t41 (cl (! (= @p_355 @p_79) :named @p_359)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_359 3 7))
(step t27.t24.t42 (cl @p_358) :rule symm :premises (t27.t24.t41))
(step t27.t24.t43 (cl @p_355) :rule resolution :premises (t27.t24.t40 t27.t24.t42 t27.a5))
(step t27.t24.t44 (cl @p_688) :rule resolution :premises (t27.t24.t35 t27.t24.t39 t27.t24.t43))
(step t27.t24.t45 (cl @p_687) :rule resolution :premises (t27.t24.t34 t27.t24.t44))
(step t27.t24.t46 (cl (=> (! (and (! (< -2 0) :named @p_680) (! (>= @p_17 1) :named @p_117)) :named @p_681) @p_679)) :rule la_mult_neg)
(step t27.t24.t47 (cl (not @p_681) @p_679) :rule implies :premises (t27.t24.t46))
(step t27.t24.t48 (cl @p_681 (not @p_680) (! (not @p_117) :named @p_119)) :rule and_neg)
(step t27.t24.t49 (cl (= (! (= @p_680 true) :named @p_686) @p_680)) :rule equiv_simplify)
(step t27.t24.t50 (cl (not @p_686) @p_680) :rule equiv1 :premises (t27.t24.t49))
(step t27.t24.t51 (cl @p_686) :rule rare_rewrite :args ("evaluate"))
(step t27.t24.t52 (cl @p_680) :rule resolution :premises (t27.t24.t50 t27.t24.t51))
(step t27.t24.t53 (cl (not (! (> @p_17 0) :named @p_682)) @p_117) :rule la_generic :args (1/1 1/1))
(step t27.t24.t54 (cl (not @p_20) (not @p_5) @p_19) :rule equiv_pos2)
(step t27.t24.t55 (cl @p_54) :rule hole :args ("TRUST_THEORY_REWRITE" @p_54 4 6))
(step t27.t24.t56 (cl (= #b0000 #b0000)) :rule refl)
(step t27.t24.t57 (cl (= s s)) :rule refl)
(step t27.t24.t58 (cl @p_53) :rule hole :args ("TRUST_THEORY_REWRITE" @p_53 4 6))
(step t27.t24.t59 (cl @p_52) :rule hole :args ("TRUST_THEORY_REWRITE" @p_52 4 6))
(step t27.t24.t60 (cl (= @p_49 @p_50)) :rule cong :premises (t27.t24.t57 t27.t24.t59))
(step t27.t24.t61 (cl @p_51) :rule hole :args ("TRUST_THEORY_REWRITE" @p_51 4 7))
(step t27.t24.t62 (cl (= @p_49 s)) :rule trans :premises (t27.t24.t60 t27.t24.t61))
(step t27.t24.t63 (cl (= @p_2 s)) :rule trans :premises (t27.t24.t58 t27.t24.t62))
(step t27.t24.t64 (cl (= @p_3 @p_45)) :rule cong :premises (t27.t24.t57 t27.t24.t63))
(step t27.t24.t65 (cl (= @p_48 @p_46)) :rule cong :premises (t27.t24.t56 t27.t24.t64))
(step t27.t24.t66 (cl (= @p_4 @p_46)) :rule trans :premises (t27.t24.t55 t27.t24.t65))
(step t27.t24.t67 (cl (= @p_5 @p_47)) :rule cong :premises (t27.t24.t66))
(step t27.t24.t68 (cl (= @p_47 @p_41)) :rule bv_intblast)
(step t27.t24.t69 (cl (= 0 0)) :rule refl)
(step t27.t24.t70 (cl @p_34) :rule refl)
(step t27.t24.t71 (cl @p_33) :rule refl)
(step t27.t24.t72 (cl @p_32) :rule refl)
(step t27.t24.t73 (cl @p_31) :rule refl)
(step t27.t24.t74 (cl @p_30) :rule refl)
(step t27.t24.t75 (cl @p_29) :rule refl)
(step t27.t24.t76 (cl @p_28) :rule refl)
(step t27.t24.t77 (cl @p_44) :rule hole :args ("TRUST_THEORY_REWRITE" @p_44 3 7))
(step t27.t24.t78 (cl (= @p_36 @p_8)) :rule cong :premises (t27.t24.t76 t27.t24.t77 t27.t24.t69))
(step t27.t24.t79 (cl (= @p_37 @p_11)) :rule cong :premises (t27.t24.t74 t27.t24.t75 t27.t24.t78))
(step t27.t24.t80 (cl (= @p_38 @p_14)) :rule cong :premises (t27.t24.t72 t27.t24.t73 t27.t24.t79))
(step t27.t24.t81 (cl (= @p_39 @p_17)) :rule cong :premises (t27.t24.t70 t27.t24.t71 t27.t24.t80))
(step t27.t24.t82 (cl (= @p_40 @p_42)) :rule cong :premises (t27.t24.t69 t27.t24.t81))
(step t27.t24.t83 (cl @p_43) :rule hole :args ("TRUST_THEORY_REWRITE" @p_43 3 7))
(step t27.t24.t84 (cl (= @p_40 @p_18)) :rule trans :premises (t27.t24.t82 t27.t24.t83))
(step t27.t24.t85 (cl (= @p_41 @p_19)) :rule cong :premises (t27.t24.t84))
(step t27.t24.t86 (cl (= @p_22 @p_21)) :rule equiv_simplify)
(step t27.t24.t87 (cl (not @p_22) @p_21) :rule equiv1 :premises (t27.t24.t86))
(step t27.t24.t88 (cl @p_23) :rule hole :args ("TRUST_THEORY_REWRITE" @p_23 3 7))
(step t27.t24.t89 (cl @p_21) :rule refl)
(step t27.t24.t90 (cl @p_27) :rule refl)
(step t27.t24.t91 (cl (= 3 3)) :rule refl)
(step t27.t24.t92 (cl @p_34) :rule cong :premises (t27.t24.t90 t27.t24.t91))
(step t27.t24.t93 (cl (= 8 8)) :rule refl)
(step t27.t24.t94 (cl @p_33) :rule cong :premises (t27.t24.t90 t27.t24.t93))
(step t27.t24.t95 (cl @p_32) :rule cong :premises (t27.t24.t90 t27.t4))
(step t27.t24.t96 (cl (= 4 4)) :rule refl)
(step t27.t24.t97 (cl @p_31) :rule cong :premises (t27.t24.t90 t27.t24.t96))
(step t27.t24.t98 (cl (= 1 1)) :rule refl)
(step t27.t24.t99 (cl @p_30) :rule cong :premises (t27.t24.t90 t27.t24.t98))
(step t27.t24.t100 (cl @p_29) :rule cong :premises (t27.t24.t90 t27.t4))
(step t27.t24.t101 (cl @p_28) :rule cong :premises (t27.t24.t90 t27.t24.t69))
(step t27.t24.t102 (cl @p_26) :rule cong :premises (t27.t24.t101 t27.t24.t90 t27.t24.t69))
(step t27.t24.t103 (cl @p_25) :rule cong :premises (t27.t24.t99 t27.t24.t100 t27.t24.t102))
(step t27.t24.t104 (cl @p_24) :rule cong :premises (t27.t24.t95 t27.t24.t97 t27.t24.t103))
(step t27.t24.t105 (cl @p_21) :rule cong :premises (t27.t24.t92 t27.t24.t94 t27.t24.t104))
(step t27.t24.t106 (cl @p_21) :rule trans :premises (t27.t24.t89 t27.t24.t105))
(step t27.t24.t107 (cl @p_23) :rule cong :premises (t27.t24.t106 t27.t24.t105))
(step t27.t24.t108 (cl @p_22) :rule hole :args ("TRUST_THEORY_REWRITE" @p_22 3 6))
(step t27.t24.t109 (cl @p_22) :rule trans :premises (t27.t24.t88 t27.t24.t107 t27.t24.t108))
(step t27.t24.t110 (cl @p_21) :rule resolution :premises (t27.t24.t87 t27.t24.t109))
(step t27.t24.t111 (cl (= @p_18 @p_18)) :rule cong :premises (t27.t24.t110 t27.t24.t69))
(step t27.t24.t112 (cl (= @p_19 @p_19)) :rule cong :premises (t27.t24.t111))
(step t27.t24.t113 (cl @p_20) :rule trans :premises (t27.t24.t67 t27.t24.t68 t27.t24.t85 t27.t24.t112))
(step t27.t24.t114 (cl @p_19) :rule resolution :premises (t27.t24.t54 t27.t24.t113 a0))
(step t27.t24.t115 (cl (! (or @p_18 (! (not (! (<= @p_17 0) :named @p_161)) :named @p_169) (! (not (! (<= 0 @p_17) :named @p_156)) :named @p_157)) :named @p_170)) :rule la_disequality)
(step t27.t24.t116 (cl @p_18 @p_169 @p_157) :rule or :premises (t27.t24.t115))
(step t27.t24.t117 (cl (! (= @p_532 @p_156) :named @p_685)) :rule comp_simplify)
(step t27.t24.t118 (cl (not @p_685) @p_535 @p_156) :rule equiv_pos2)
(step t27.t24.t119 (cl @p_156) :rule resolution :premises (t27.t24.t117 t27.t24.t118 t27.a4))
(step t27.t24.t120 (cl (! (= @p_682 @p_169) :named @p_684)) :rule comp_simplify)
(step t27.t24.t121 (cl (not @p_684) @p_682 (! (not @p_169) :named @p_683)) :rule equiv_pos1)
(step t27.t24.t122 (cl @p_682 @p_683) :rule resolution :premises (t27.t24.t120 t27.t24.t121))
(step t27.t24.t123 (cl @p_682) :rule resolution :premises (t27.t24.t114 t27.t24.t116 t27.t24.t119 t27.t24.t122))
(step t27.t24.t124 (cl @p_117) :rule resolution :premises (t27.t24.t53 t27.t24.t123))
(step t27.t24.t125 (cl @p_681) :rule resolution :premises (t27.t24.t48 t27.t24.t52 t27.t24.t124))
(step t27.t24.t126 (cl @p_679) :rule resolution :premises (t27.t24.t47 t27.t24.t125))
(step t27.t24.t127 (cl (! (not (! (= @p_515 @p_518) :named @p_528)) :named @p_530) @p_517 @p_518) :rule equiv_pos2)
(step t27.t24.t128 (cl (! (= @p_518 @p_515) :named @p_529)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_529 3 6))
(step t27.t24.t129 (cl @p_528) :rule symm :premises (t27.t24.t128))
(step t27.t24.t130 (cl @p_518) :rule resolution :premises (t27.t24.t127 t27.t24.t129 t27.a2))
(step t27.t24.t131 (cl (=> (! (and @p_672 (! (= @p_558 0) :named @p_608)) :named @p_677) @p_676)) :rule la_mult_pos)
(step t27.t24.t132 (cl (not @p_677) @p_676) :rule implies :premises (t27.t24.t131))
(step t27.t24.t133 (cl @p_677 @p_675 (! (not @p_608) :named @p_613)) :rule and_neg)
(step t27.t24.t134 (cl (! (not (! (= @p_252 @p_608) :named @p_610)) :named @p_612) @p_486 @p_608) :rule equiv_pos2)
(step t27.t24.t135 (cl (! (= @p_608 @p_252) :named @p_611)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_611 3 7))
(step t27.t24.t136 (cl @p_610) :rule symm :premises (t27.t24.t135))
(step t27.t24.t137 (cl @p_608) :rule resolution :premises (t27.t24.t134 t27.t24.t136 t27.a1))
(step t27.t24.t138 (cl @p_677) :rule resolution :premises (t27.t24.t133 t27.t24.t39 t27.t24.t137))
(step t27.t24.t139 (cl @p_676) :rule resolution :premises (t27.t24.t132 t27.t24.t138))
(step t27.t24.t140 (cl (=> (! (and @p_672 (! (= @p_548 0) :named @p_566)) :named @p_673) @p_671)) :rule la_mult_pos)
(step t27.t24.t141 (cl (not @p_673) @p_671) :rule implies :premises (t27.t24.t140))
(step t27.t24.t142 (cl @p_673 @p_675 (! (not @p_566) :named @p_572)) :rule and_neg)
(step t27.t24.t143 (cl (! (not (! (= @p_433 @p_566) :named @p_569)) :named @p_571) @p_533 @p_566) :rule equiv_pos2)
(step t27.t24.t144 (cl (! (= @p_566 @p_433) :named @p_570)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_570 3 7))
(step t27.t24.t145 (cl @p_569) :rule symm :premises (t27.t24.t144))
(step t27.t24.t146 (cl @p_566) :rule resolution :premises (t27.t24.t143 t27.t24.t145 t27.a0))
(step t27.t24.t147 (cl @p_673) :rule resolution :premises (t27.t24.t142 t27.t24.t39 t27.t24.t146))
(step t27.t24.t148 (cl @p_671) :rule resolution :premises (t27.t24.t141 t27.t24.t147))
(step t27.t24.t149 (cl @p_670) :rule resolution :premises (t27.t24.t23 t27.t24.t32 t27.t24.t45 t27.t24.t126 t27.t24.t130 t27.t24.t139 t27.t24.t148))
(step t27.t24.t150 (cl false) :rule resolution :premises (t27.t24.t0 t27.t24.t22 t27.t24.t149))
(step t27.t24 (cl @p_650 false) :rule subproof :discharge (t27.t24.a0))
(step t27.t25 (cl @p_660 false) :rule resolution :premises (t27.t23 t27.t24))
(step t27.t26 (cl @p_660 (! (not false) :named @p_132)) :rule implies_neg2)
(step t27.t27 (cl @p_660 @p_660) :rule resolution :premises (t27.t25 t27.t26))
(step t27.t28 (cl @p_660) :rule contraction :premises (t27.t27))
(step t27.t29 (cl (= @p_660 @p_650)) :rule implies_simplify)
(step t27.t30 (cl (not @p_660) @p_650) :rule equiv1 :premises (t27.t29))
(step t27.t31 (cl @p_650) :rule resolution :premises (t27.t28 t27.t30))
(step t27.t32 (cl @p_659) :rule resolution :premises (t27.t21 t27.t22 t27.t31))
(step t27.t33 (cl @p_650) :rule resolution :premises (t27.t19 t27.t20 t27.t32))
(step t27.t34 (cl) :rule resolution :premises (t27.t18 t27.t33))
(step t27 (cl @p_533 @p_486 @p_517 @p_122 @p_535 @p_312 @p_650 false) :rule subproof :discharge (t27.a0 t27.a1 t27.a2 t27.a3 t27.a4 t27.a5 t27.a6))
(step t28 (cl @p_657 @p_433) :rule and_pos :args (0))
(step t29 (cl @p_657 @p_252) :rule and_pos :args (1))
(step t30 (cl @p_657 @p_515) :rule and_pos :args (2))
(step t31 (cl @p_657 @p_19) :rule and_pos :args (3))
(step t32 (cl @p_657 @p_532) :rule and_pos :args (4))
(step t33 (cl @p_657 @p_79) :rule and_pos :args (5))
(step t34 (cl @p_657 @p_57) :rule and_pos :args (6))
(step t35 (cl false @p_657 @p_657 @p_657 @p_657 @p_657 @p_657 @p_657) :rule resolution :premises (t27 t28 t29 t30 t31 t32 t33 t34))
(step t36 (cl @p_657 @p_657 @p_657 @p_657 @p_657 @p_657 @p_657 false) :rule reordering :premises (t35))
(step t37 (cl @p_657 false) :rule contraction :premises (t36))
(step t38 (cl @p_658 false) :rule resolution :premises (t26 t37))
(step t39 (cl @p_658 @p_132) :rule implies_neg2)
(step t40 (cl @p_658 @p_658) :rule resolution :premises (t38 t39))
(step t41 (cl @p_658) :rule contraction :premises (t40))
(step t42 (cl (= @p_658 @p_657)) :rule implies_simplify)
(step t43 (cl (not @p_658) @p_657) :rule equiv1 :premises (t42))
(step t44 (cl @p_657) :rule resolution :premises (t41 t43))
(step t45 (cl @p_655) :rule resolution :premises (t25 t44))
(step t46 (cl @p_122 @p_312 @p_486 @p_533 @p_650 @p_535 @p_517) :rule not_and :premises (t45))
(step t47 (cl @p_652 (! (not @p_122) :named @p_125)) :rule or_neg :args (0))
(step t48 (cl @p_652 (! (not @p_312) :named @p_318)) :rule or_neg :args (1))
(step t49 (cl @p_652 (! (not @p_486) :named @p_490)) :rule or_neg :args (2))
(step t50 (cl @p_652 (! (not @p_533) :named @p_539)) :rule or_neg :args (3))
(step t51 (cl @p_652 @p_653) :rule or_neg :args (4))
(step t52 (cl @p_652 (! (not @p_535) :named @p_536)) :rule or_neg :args (5))
(step t53 (cl @p_652 (not @p_517)) :rule or_neg :args (6))
(step t54 (cl @p_652 @p_652 @p_652 @p_652 @p_652 @p_652 @p_652) :rule resolution :premises (t46 t47 t48 t49 t50 t51 t52 t53))
(step t55 (cl @p_652) :rule contraction :premises (t54))
(step t56 (cl @p_651) :rule resolution :premises (t0 t8 t55))
(step t57 (cl @p_18 @p_312 @p_486 @p_533 @p_650 @p_535 @p_514) :rule or :premises (t56))
(step t58 (cl (not (! (= (! (or @p_486 @p_312 (! (not @p_10) :named @p_449) (! (not @p_59) :named @p_537) @p_533 @p_536) :named @p_538) (! (or @p_486 @p_312 @p_449 @p_58 @p_533 @p_532) :named @p_534)) :named @p_647)) (not @p_538) @p_534) :rule equiv_pos2)
(step t59 (cl (= @p_449 @p_449)) :rule refl)
(step t60 (cl (! (= @p_537 @p_58) :named @p_649)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_649 1 6))
(step t61 (cl (! (= @p_536 @p_532) :named @p_648)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_648 1 6))
(step t62 (cl @p_647) :rule cong :premises (t3 t2 t59 t60 t4 t61))
(step t63 (cl (! (=> (! (and @p_252 @p_79 @p_10 @p_59 @p_433 @p_535) :named @p_540) false) :named @p_542) @p_540) :rule implies_neg1)
(anchor :step t64)
(assume t64.a0 @p_252)
(assume t64.a1 @p_79)
(assume t64.a2 @p_10)
(assume t64.a3 @p_59)
(assume t64.a4 @p_433)
(assume t64.a5 @p_535)
(step t64.t0 (cl (not (! (= (! (< (! (+ (! (to_real @p_17) :named @p_561) (! (to_real (! (* -1 @p_558) :named @p_559)) :named @p_560) (! (to_real (! (* -1 @p_334) :named @p_555)) :named @p_556) (! (* -1/2 (! (to_real @p_6) :named @p_553)) :named @p_554) (! (* 1/2 (! (to_real @p_56) :named @p_551)) :named @p_552) (! (to_real (! (* -1 @p_548) :named @p_549)) :named @p_550)) :named @p_562) (! (+ 0/1 (! (to_real @p_136) :named @p_543) @p_543 (! (* -1/2 1/1) :named @p_545) (! (* 1/2 1/1) :named @p_544) @p_543) :named @p_546)) :named @p_563) false) :named @p_624)) (not @p_563) false) :rule equiv_pos2)
(step t64.t1 (cl (! (= @p_563 (! (not (! (>= @p_562 @p_546) :named @p_625)) :named @p_626)) :named @p_646)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_646 3 6))
(step t64.t2 (cl (! (= @p_561 @p_561) :named @p_622)) :rule refl)
(step t64.t3 (cl (! (= @p_559 (! (+ @p_333 @p_11) :named @p_638)) :named @p_645)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_645 3 7))
(step t64.t4 (cl (= @p_560 (! (to_real @p_638) :named @p_639))) :rule cong :premises (t64.t3))
(step t64.t5 (cl (! (= @p_555 (! (+ (* -1 @p_17) @p_14) :named @p_636)) :named @p_644)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_644 3 7))
(step t64.t6 (cl (= @p_556 (! (to_real @p_636) :named @p_637))) :rule cong :premises (t64.t5))
(step t64.t7 (cl (= @p_554 @p_554)) :rule refl)
(step t64.t8 (cl (! (= @p_552 (! (+ (* 1/2 @p_553) (to_real @p_547)) :named @p_635)) :named @p_643)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_643 3 7))
(step t64.t9 (cl (! (= @p_549 (! (+ @p_557 @p_9) :named @p_633)) :named @p_642)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_642 3 7))
(step t64.t10 (cl (= @p_550 (! (to_real @p_633) :named @p_634))) :rule cong :premises (t64.t9))
(step t64.t11 (cl (= @p_562 (! (+ @p_561 @p_639 @p_637 @p_554 @p_635 @p_634) :named @p_640))) :rule cong :premises (t64.t2 t64.t4 t64.t6 t64.t7 t64.t8 t64.t10))
(step t64.t12 (cl (! (= @p_640 0/1) :named @p_641)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_641 3 7))
(step t64.t13 (cl (= @p_562 0/1)) :rule trans :premises (t64.t11 t64.t12))
(step t64.t14 (cl (= 0/1 0/1)) :rule refl)
(step t64.t15 (cl @p_199) :rule hole :args ("TRUST_THEORY_REWRITE" @p_199 3 6))
(step t64.t16 (cl (= @p_543 (! (to_real 0) :named @p_619))) :rule cong :premises (t64.t15))
(step t64.t17 (cl (! (= @p_619 0/1) :named @p_621)) :rule rare_rewrite :args ("evaluate"))
(step t64.t18 (cl (= @p_543 0/1)) :rule trans :premises (t64.t16 t64.t17))
(step t64.t19 (cl (! (= @p_545 -1/2) :named @p_632)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_632 3 7))
(step t64.t20 (cl (! (= @p_544 1/2) :named @p_631)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_631 3 7))
(step t64.t21 (cl (= @p_546 (! (+ 0/1 0/1 0/1 -1/2 1/2 0/1) :named @p_629))) :rule cong :premises (t64.t14 t64.t18 t64.t18 t64.t19 t64.t20 t64.t18))
(step t64.t22 (cl (! (= @p_629 0/1) :named @p_630)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_630 3 7))
(step t64.t23 (cl (= @p_546 0/1)) :rule trans :premises (t64.t21 t64.t22))
(step t64.t24 (cl (= @p_625 (! (>= 0/1 0/1) :named @p_627))) :rule cong :premises (t64.t13 t64.t23))
(step t64.t25 (cl (! (= @p_627 true) :named @p_628)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_628 3 7))
(step t64.t26 (cl (= @p_625 true)) :rule trans :premises (t64.t24 t64.t25))
(step t64.t27 (cl (= @p_626 @p_193)) :rule cong :premises (t64.t26))
(step t64.t28 (cl @p_194) :rule hole :args ("TRUST_THEORY_REWRITE" @p_194 1 7))
(step t64.t29 (cl (= @p_626 false)) :rule trans :premises (t64.t27 t64.t28))
(step t64.t30 (cl @p_624) :rule trans :premises (t64.t1 t64.t29))
(step t64.t31 (cl (not (! (< @p_561 0/1) :named @p_614)) (not (! (= @p_560 @p_543) :named @p_606)) (not (! (= @p_556 @p_543) :named @p_603)) (! (not (! (= @p_554 @p_545) :named @p_593)) :named @p_600) (! (not (! (<= @p_552 @p_544) :named @p_573)) :named @p_585) (not (! (= @p_550 @p_543) :named @p_564)) @p_563) :rule la_generic :args (1/1 -1/1 -1/1 -1/1 1/1 -1/1 1/1))
(step t64.t32 (cl (not (! (= (! (< @p_17 0) :named @p_155) @p_614) :named @p_618)) (! (not @p_155) :named @p_171) @p_614) :rule equiv_pos2)
(step t64.t33 (cl (! (= @p_155 (! (< @p_561 @p_619) :named @p_620)) :named @p_623)) :rule hole :args ("ARITH_PRED_CAST_TYPE" "UNKNOWN_THEORY" @p_623 @p_623))
(step t64.t34 (cl @p_622) :rule refl)
(step t64.t35 (cl @p_621) :rule rare_rewrite :args ("evaluate"))
(step t64.t36 (cl (= @p_620 @p_614)) :rule cong :premises (t64.t34 t64.t35))
(step t64.t37 (cl @p_618) :rule trans :premises (t64.t33 t64.t36))
(step t64.t38 (cl (not (! (= @p_535 @p_155) :named @p_616)) @p_536 @p_155) :rule equiv_pos2)
(step t64.t39 (cl (! (= @p_155 @p_535) :named @p_617)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_617 3 6))
(step t64.t40 (cl @p_616) :rule symm :premises (t64.t39))
(step t64.t41 (cl @p_155) :rule resolution :premises (t64.t38 t64.t40 t64.a5))
(step t64.t42 (cl @p_614) :rule resolution :premises (t64.t32 t64.t37 t64.t41))
(step t64.t43 (cl (=> (! (and @p_147 @p_608) :named @p_609) (! (= @p_559 @p_136) :named @p_607))) :rule la_mult_neg)
(step t64.t44 (cl (not @p_609) @p_607) :rule implies :premises (t64.t43))
(step t64.t45 (cl @p_609 @p_150 @p_613) :rule and_neg)
(step t64.t46 (cl (= @p_149 @p_147)) :rule equiv_simplify)
(step t64.t47 (cl (not @p_149) @p_147) :rule equiv1 :premises (t64.t46))
(step t64.t48 (cl @p_149) :rule rare_rewrite :args ("evaluate"))
(step t64.t49 (cl @p_147) :rule resolution :premises (t64.t47 t64.t48))
(step t64.t50 (cl @p_612 @p_486 @p_608) :rule equiv_pos2)
(step t64.t51 (cl @p_611) :rule hole :args ("TRUST_THEORY_REWRITE" @p_611 3 7))
(step t64.t52 (cl @p_610) :rule symm :premises (t64.t51))
(step t64.t53 (cl @p_608) :rule resolution :premises (t64.t50 t64.t52 t64.a0))
(step t64.t54 (cl @p_609) :rule resolution :premises (t64.t45 t64.t49 t64.t53))
(step t64.t55 (cl @p_607) :rule resolution :premises (t64.t44 t64.t54))
(step t64.t56 (cl @p_606) :rule cong :premises (t64.t55))
(step t64.t57 (cl (=> (! (and @p_147 @p_355) :named @p_605) (! (= @p_555 @p_136) :named @p_604))) :rule la_mult_neg)
(step t64.t58 (cl (not @p_605) @p_604) :rule implies :premises (t64.t57))
(step t64.t59 (cl @p_605 @p_150 @p_361) :rule and_neg)
(step t64.t60 (cl @p_360 @p_312 @p_355) :rule equiv_pos2)
(step t64.t61 (cl @p_359) :rule hole :args ("TRUST_THEORY_REWRITE" @p_359 3 7))
(step t64.t62 (cl @p_358) :rule symm :premises (t64.t61))
(step t64.t63 (cl @p_355) :rule resolution :premises (t64.t60 t64.t62 t64.a1))
(step t64.t64 (cl @p_605) :rule resolution :premises (t64.t59 t64.t49 t64.t63))
(step t64.t65 (cl @p_604) :rule resolution :premises (t64.t58 t64.t64))
(step t64.t66 (cl @p_603) :rule cong :premises (t64.t65))
(step t64.t67 (cl (! (=> (! (and (! (< -1/2 0/1) :named @p_594) @p_10) :named @p_595) @p_593) :named @p_599) @p_595) :rule implies_neg1)
(anchor :step t64.t68)
(assume t64.t68.a0 @p_594)
(assume t64.t68.a1 @p_10)
(step t64.t68.t0 (cl (=> (! (and @p_594 (! (= @p_553 1/1) :named @p_601)) :named @p_602) @p_593)) :rule la_mult_neg)
(step t64.t68.t1 (cl (not @p_602) @p_593) :rule implies :premises (t64.t68.t0))
(step t64.t68.t2 (cl @p_602 (! (not @p_594) :named @p_597) (not @p_601)) :rule and_neg)
(step t64.t68.t3 (cl (= @p_553 (! (to_real 1) :named @p_589))) :rule cong :premises (t64.t68.a1))
(step t64.t68.t4 (cl (! (= @p_589 1/1) :named @p_591)) :rule rare_rewrite :args ("evaluate"))
(step t64.t68.t5 (cl @p_601) :rule trans :premises (t64.t68.t3 t64.t68.t4))
(step t64.t68.t6 (cl @p_602) :rule resolution :premises (t64.t68.t2 t64.t68.a0 t64.t68.t5))
(step t64.t68.t7 (cl @p_593) :rule resolution :premises (t64.t68.t1 t64.t68.t6))
(step t64.t68 (cl @p_597 @p_449 @p_593) :rule subproof :discharge (t64.t68.a0 t64.t68.a1))
(step t64.t69 (cl (! (not @p_595) :named @p_598) @p_594) :rule and_pos :args (0))
(step t64.t70 (cl @p_598 @p_10) :rule and_pos :args (1))
(step t64.t71 (cl @p_593 @p_598 @p_598) :rule resolution :premises (t64.t68 t64.t69 t64.t70))
(step t64.t72 (cl @p_598 @p_598 @p_593) :rule reordering :premises (t64.t71))
(step t64.t73 (cl @p_598 @p_593) :rule contraction :premises (t64.t72))
(step t64.t74 (cl @p_599 @p_593) :rule resolution :premises (t64.t67 t64.t73))
(step t64.t75 (cl @p_599 @p_600) :rule implies_neg2)
(step t64.t76 (cl @p_599 @p_599) :rule resolution :premises (t64.t74 t64.t75))
(step t64.t77 (cl @p_599) :rule contraction :premises (t64.t76))
(step t64.t78 (cl @p_598 @p_593) :rule implies :premises (t64.t77))
(step t64.t79 (cl @p_595 @p_597 @p_449) :rule and_neg)
(step t64.t80 (cl (= (! (= @p_594 true) :named @p_596) @p_594)) :rule equiv_simplify)
(step t64.t81 (cl (not @p_596) @p_594) :rule equiv1 :premises (t64.t80))
(step t64.t82 (cl @p_596) :rule rare_rewrite :args ("evaluate"))
(step t64.t83 (cl @p_594) :rule resolution :premises (t64.t81 t64.t82))
(step t64.t84 (cl @p_595) :rule resolution :premises (t64.t79 t64.t83 t64.a2))
(step t64.t85 (cl @p_593) :rule resolution :premises (t64.t78 t64.t84))
(step t64.t86 (cl (! (=> (! (and (! (> 1/2 0/1) :named @p_575) (! (<= @p_56 1) :named @p_574)) :named @p_576) @p_573) :named @p_584) @p_576) :rule implies_neg1)
(anchor :step t64.t87)
(assume t64.t87.a0 @p_575)
(assume t64.t87.a1 @p_574)
(step t64.t87.t0 (cl (=> (! (and @p_575 (! (<= @p_551 1/1) :named @p_586)) :named @p_587) @p_573)) :rule la_mult_pos)
(step t64.t87.t1 (cl (not @p_587) @p_573) :rule implies :premises (t64.t87.t0))
(step t64.t87.t2 (cl @p_587 (! (not @p_575) :named @p_581) (not @p_586)) :rule and_neg)
(step t64.t87.t3 (cl (not (! (= @p_574 @p_586) :named @p_588)) (! (not @p_574) :named @p_582) @p_586) :rule equiv_pos2)
(step t64.t87.t4 (cl (! (= @p_574 (! (<= @p_551 @p_589) :named @p_590)) :named @p_592)) :rule hole :args ("ARITH_PRED_CAST_TYPE" "UNKNOWN_THEORY" @p_592 @p_592))
(step t64.t87.t5 (cl (= @p_551 @p_551)) :rule refl)
(step t64.t87.t6 (cl @p_591) :rule rare_rewrite :args ("evaluate"))
(step t64.t87.t7 (cl (= @p_590 @p_586)) :rule cong :premises (t64.t87.t5 t64.t87.t6))
(step t64.t87.t8 (cl @p_588) :rule trans :premises (t64.t87.t4 t64.t87.t7))
(step t64.t87.t9 (cl @p_586) :rule resolution :premises (t64.t87.t3 t64.t87.t8 t64.t87.a1))
(step t64.t87.t10 (cl @p_587) :rule resolution :premises (t64.t87.t2 t64.t87.a0 t64.t87.t9))
(step t64.t87.t11 (cl @p_573) :rule resolution :premises (t64.t87.t1 t64.t87.t10))
(step t64.t87 (cl @p_581 @p_582 @p_573) :rule subproof :discharge (t64.t87.a0 t64.t87.a1))
(step t64.t88 (cl (! (not @p_576) :named @p_583) @p_575) :rule and_pos :args (0))
(step t64.t89 (cl @p_583 @p_574) :rule and_pos :args (1))
(step t64.t90 (cl @p_573 @p_583 @p_583) :rule resolution :premises (t64.t87 t64.t88 t64.t89))
(step t64.t91 (cl @p_583 @p_583 @p_573) :rule reordering :premises (t64.t90))
(step t64.t92 (cl @p_583 @p_573) :rule contraction :premises (t64.t91))
(step t64.t93 (cl @p_584 @p_573) :rule resolution :premises (t64.t86 t64.t92))
(step t64.t94 (cl @p_584 @p_585) :rule implies_neg2)
(step t64.t95 (cl @p_584 @p_584) :rule resolution :premises (t64.t93 t64.t94))
(step t64.t96 (cl @p_584) :rule contraction :premises (t64.t95))
(step t64.t97 (cl @p_583 @p_573) :rule implies :premises (t64.t96))
(step t64.t98 (cl @p_576 @p_581 @p_582) :rule and_neg)
(step t64.t99 (cl (= (! (= @p_575 true) :named @p_580) @p_575)) :rule equiv_simplify)
(step t64.t100 (cl (not @p_580) @p_575) :rule equiv1 :premises (t64.t99))
(step t64.t101 (cl @p_580) :rule rare_rewrite :args ("evaluate"))
(step t64.t102 (cl @p_575) :rule resolution :premises (t64.t100 t64.t101))
(step t64.t103 (cl (not (! (< @p_56 2) :named @p_577)) @p_574) :rule la_generic :args (1/1 1/1))
(step t64.t104 (cl (not (! (= @p_59 @p_577) :named @p_578)) @p_537 @p_577) :rule equiv_pos2)
(step t64.t105 (cl (! (= @p_577 @p_59) :named @p_579)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_579 3 6))
(step t64.t106 (cl @p_578) :rule symm :premises (t64.t105))
(step t64.t107 (cl (not @p_67) (not @p_66) @p_60) :rule equiv_pos2)
(step t64.t108 (cl @p_78) :rule hole :args ("TRUST_THEORY_REWRITE" @p_78 3 7))
(step t64.t109 (cl @p_77) :rule hole :args ("TRUST_THEORY_REWRITE" @p_77 3 6))
(step t64.t110 (cl @p_27) :rule refl)
(step t64.t111 (cl (= 2 2)) :rule refl)
(step t64.t112 (cl @p_76) :rule hole :args ("TRUST_THEORY_REWRITE" @p_76 3 7))
(step t64.t113 (cl (= @p_62 @p_74)) :rule cong :premises (t64.t111 t64.t112))
(step t64.t114 (cl @p_75) :rule hole :args ("TRUST_THEORY_REWRITE" @p_75 3 7))
(step t64.t115 (cl (= @p_62 @p_70)) :rule trans :premises (t64.t113 t64.t114))
(step t64.t116 (cl (= @p_68 @p_71)) :rule cong :premises (t64.t110 t64.t115))
(step t64.t117 (cl @p_72) :rule hole :args ("TRUST_THEORY_REWRITE" @p_72 3 7))
(step t64.t118 (cl (= @p_68 @p_58)) :rule trans :premises (t64.t116 t64.t117))
(step t64.t119 (cl (= @p_69 @p_59)) :rule cong :premises (t64.t118))
(step t64.t120 (cl (= @p_63 @p_59)) :rule trans :premises (t64.t109 t64.t119))
(step t64.t121 (cl @p_67) :rule cong :premises (t64.t108 t64.t120))
(step t64.t122 (cl (and @p_29 @p_66)) :rule hole :args ("ARITH_REDUCTION" @p_9))
(step t64.t123 (cl @p_66) :rule and :premises (t64.t122) :args (1))
(step t64.t124 (cl @p_60) :rule resolution :premises (t64.t107 t64.t121 t64.t123))
(step t64.t125 (cl @p_59) :rule and :premises (t64.t124) :args (1))
(step t64.t126 (cl @p_577) :rule resolution :premises (t64.t104 t64.t106 t64.t125))
(step t64.t127 (cl @p_574) :rule resolution :premises (t64.t103 t64.t126))
(step t64.t128 (cl @p_576) :rule resolution :premises (t64.t98 t64.t102 t64.t127))
(step t64.t129 (cl @p_573) :rule resolution :premises (t64.t97 t64.t128))
(step t64.t130 (cl (=> (! (and @p_147 @p_566) :named @p_567) (! (= @p_549 @p_136) :named @p_565))) :rule la_mult_neg)
(step t64.t131 (cl (not @p_567) @p_565) :rule implies :premises (t64.t130))
(step t64.t132 (cl @p_567 @p_150 @p_572) :rule and_neg)
(step t64.t133 (cl @p_571 @p_533 @p_566) :rule equiv_pos2)
(step t64.t134 (cl @p_570) :rule hole :args ("TRUST_THEORY_REWRITE" @p_570 3 7))
(step t64.t135 (cl @p_569) :rule symm :premises (t64.t134))
(step t64.t136 (cl @p_566) :rule resolution :premises (t64.t133 t64.t135 t64.a4))
(step t64.t137 (cl @p_567) :rule resolution :premises (t64.t132 t64.t49 t64.t136))
(step t64.t138 (cl @p_565) :rule resolution :premises (t64.t131 t64.t137))
(step t64.t139 (cl @p_564) :rule cong :premises (t64.t138))
(step t64.t140 (cl @p_563) :rule resolution :premises (t64.t31 t64.t42 t64.t56 t64.t66 t64.t85 t64.t129 t64.t139))
(step t64.t141 (cl false) :rule resolution :premises (t64.t0 t64.t30 t64.t140))
(step t64 (cl @p_486 @p_312 @p_449 @p_537 @p_533 @p_536 false) :rule subproof :discharge (t64.a0 t64.a1 t64.a2 t64.a3 t64.a4 t64.a5))
(step t65 (cl (! (not @p_540) :named @p_541) @p_252) :rule and_pos :args (0))
(step t66 (cl @p_541 @p_79) :rule and_pos :args (1))
(step t67 (cl @p_541 @p_10) :rule and_pos :args (2))
(step t68 (cl @p_541 @p_59) :rule and_pos :args (3))
(step t69 (cl @p_541 @p_433) :rule and_pos :args (4))
(step t70 (cl @p_541 @p_535) :rule and_pos :args (5))
(step t71 (cl false @p_541 @p_541 @p_541 @p_541 @p_541 @p_541) :rule resolution :premises (t64 t65 t66 t67 t68 t69 t70))
(step t72 (cl @p_541 @p_541 @p_541 @p_541 @p_541 @p_541 false) :rule reordering :premises (t71))
(step t73 (cl @p_541 false) :rule contraction :premises (t72))
(step t74 (cl @p_542 false) :rule resolution :premises (t63 t73))
(step t75 (cl @p_542 @p_132) :rule implies_neg2)
(step t76 (cl @p_542 @p_542) :rule resolution :premises (t74 t75))
(step t77 (cl @p_542) :rule contraction :premises (t76))
(step t78 (cl (= @p_542 @p_541)) :rule implies_simplify)
(step t79 (cl (not @p_542) @p_541) :rule equiv1 :premises (t78))
(step t80 (cl @p_541) :rule resolution :premises (t77 t79))
(step t81 (cl @p_486 @p_312 @p_449 @p_537 @p_533 @p_536) :rule not_and :premises (t80))
(step t82 (cl @p_538 @p_490) :rule or_neg :args (0))
(step t83 (cl @p_538 @p_318) :rule or_neg :args (1))
(step t84 (cl @p_538 (not @p_449)) :rule or_neg :args (2))
(step t85 (cl @p_538 (not @p_537)) :rule or_neg :args (3))
(step t86 (cl @p_538 @p_539) :rule or_neg :args (4))
(step t87 (cl @p_538 (not @p_536)) :rule or_neg :args (5))
(step t88 (cl @p_538 @p_538 @p_538 @p_538 @p_538 @p_538) :rule resolution :premises (t81 t82 t83 t84 t85 t86 t87))
(step t89 (cl @p_538) :rule contraction :premises (t88))
(step t90 (cl @p_534) :rule resolution :premises (t58 t62 t89))
(step t91 (cl @p_486 @p_312 @p_449 @p_58 @p_533 @p_532) :rule or :premises (t90))
(step t92 (cl @p_58 @p_312 @p_486 @p_449 @p_533 @p_532) :rule reordering :premises (t91))
(step t93 (cl (not @p_67) (not @p_66) @p_60) :rule equiv_pos2)
(step t94 (cl @p_78) :rule hole :args ("TRUST_THEORY_REWRITE" @p_78 3 7))
(step t95 (cl @p_77) :rule hole :args ("TRUST_THEORY_REWRITE" @p_77 3 6))
(step t96 (cl @p_27) :rule refl)
(step t97 (cl (= 2 2)) :rule refl)
(step t98 (cl @p_76) :rule hole :args ("TRUST_THEORY_REWRITE" @p_76 3 7))
(step t99 (cl (= @p_62 @p_74)) :rule cong :premises (t97 t98))
(step t100 (cl @p_75) :rule hole :args ("TRUST_THEORY_REWRITE" @p_75 3 7))
(step t101 (cl (= @p_62 @p_70)) :rule trans :premises (t99 t100))
(step t102 (cl (= @p_68 @p_71)) :rule cong :premises (t96 t101))
(step t103 (cl @p_72) :rule hole :args ("TRUST_THEORY_REWRITE" @p_72 3 7))
(step t104 (cl (= @p_68 @p_58)) :rule trans :premises (t102 t103))
(step t105 (cl (= @p_69 @p_59)) :rule cong :premises (t104))
(step t106 (cl (= @p_63 @p_59)) :rule trans :premises (t95 t105))
(step t107 (cl @p_67) :rule cong :premises (t94 t106))
(step t108 (cl (and @p_29 @p_66)) :rule hole :args ("ARITH_REDUCTION" @p_9))
(step t109 (cl @p_66) :rule and :premises (t108) :args (1))
(step t110 (cl @p_60) :rule resolution :premises (t93 t107 t109))
(step t111 (cl @p_59) :rule and :premises (t110) :args (1))
(step t112 (cl (! (not (! (= (! (ite @p_16 (! (= @p_17 @p_15) :named @p_81) @p_79) :named @p_82) @p_82) :named @p_83)) :named @p_86) (! (not @p_82) :named @p_87) @p_82) :rule equiv_pos2)
(step t113 (cl @p_34) :rule refl)
(step t114 (cl @p_21) :rule refl)
(step t115 (cl (and @p_33 (! (and (! (<= (! (* 8 @p_15) :named @p_93) @p_6) :named @p_94) (! (< @p_6 (! (* 8 (! (+ @p_15 1) :named @p_90)) :named @p_91)) :named @p_92)) :named @p_95))) :rule hole :args ("ARITH_REDUCTION" @p_15))
(step t116 (cl @p_33) :rule and :premises (t115) :args (0))
(step t117 (cl (! (= @p_81 @p_81) :named @p_85)) :rule cong :premises (t114 t116))
(step t118 (cl (= (! (= @p_24 true) :named @p_88) @p_24)) :rule equiv_simplify)
(step t119 (cl (not @p_88) @p_24) :rule equiv1 :premises (t118))
(step t120 (cl (! (= @p_24 @p_24) :named @p_89)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_89 3 7))
(step t121 (cl @p_24) :rule refl)
(step t122 (cl @p_27) :rule refl)
(step t123 (cl @p_32) :rule cong :premises (t122 t97))
(step t124 (cl (= 4 4)) :rule refl)
(step t125 (cl @p_31) :rule cong :premises (t122 t124))
(step t126 (cl (= 1 1)) :rule refl)
(step t127 (cl @p_30) :rule cong :premises (t122 t126))
(step t128 (cl @p_29) :rule cong :premises (t122 t97))
(step t129 (cl (= 0 0)) :rule refl)
(step t130 (cl @p_28) :rule cong :premises (t122 t129))
(step t131 (cl @p_26) :rule cong :premises (t130 t122 t129))
(step t132 (cl @p_25) :rule cong :premises (t127 t128 t131))
(step t133 (cl @p_24) :rule cong :premises (t123 t125 t132))
(step t134 (cl @p_24) :rule trans :premises (t121 t133))
(step t135 (cl @p_89) :rule cong :premises (t134 t133))
(step t136 (cl @p_88) :rule hole :args ("TRUST_THEORY_REWRITE" @p_88 3 6))
(step t137 (cl @p_88) :rule trans :premises (t120 t135 t136))
(step t138 (cl @p_24) :rule resolution :premises (t119 t137))
(step t139 (cl (! (= @p_79 @p_79) :named @p_84)) :rule cong :premises (t114 t138))
(step t140 (cl @p_83) :rule cong :premises (t113 t117 t139))
(step t141 (cl @p_86 @p_87 @p_82) :rule equiv_pos2)
(step t142 (cl (= @p_22 @p_21)) :rule equiv_simplify)
(step t143 (cl (not @p_22) @p_21) :rule equiv1 :premises (t142))
(step t144 (cl @p_23) :rule hole :args ("TRUST_THEORY_REWRITE" @p_23 3 7))
(step t145 (cl @p_21) :rule refl)
(step t146 (cl (= 3 3)) :rule refl)
(step t147 (cl @p_34) :rule cong :premises (t122 t146))
(step t148 (cl (= 8 8)) :rule refl)
(step t149 (cl @p_33) :rule cong :premises (t122 t148))
(step t150 (cl @p_21) :rule cong :premises (t147 t149 t133))
(step t151 (cl @p_21) :rule trans :premises (t145 t150))
(step t152 (cl @p_23) :rule cong :premises (t151 t150))
(step t153 (cl @p_22) :rule hole :args ("TRUST_THEORY_REWRITE" @p_22 3 6))
(step t154 (cl @p_22) :rule trans :premises (t144 t152 t153))
(step t155 (cl @p_21) :rule resolution :premises (t143 t154))
(step t156 (cl @p_33) :rule refl)
(step t157 (cl @p_85) :rule cong :premises (t155 t156))
(step t158 (cl @p_24) :rule refl)
(step t159 (cl @p_84) :rule cong :premises (t155 t158))
(step t160 (cl @p_83) :rule cong :premises (t113 t157 t159))
(step t161 (cl @p_82) :rule hole)
(step t162 (cl @p_82) :rule resolution :premises (t141 t160 t161))
(step t163 (cl @p_82) :rule resolution :premises (t112 t140 t162))
(step t164 (cl @p_16 @p_79) :rule ite1 :premises (t163))
(step t165 (cl (! (=> (! (and @p_16 @p_81 (! (>= (! (+ @p_6 (! (* -8 @p_15) :named @p_97)) :named @p_98) 0) :named @p_101) @p_117) :named @p_215) (! (and @p_16 @p_117 @p_81 @p_101) :named @p_217)) :named @p_251) @p_215) :rule implies_neg1)
(anchor :step t166)
(assume t166.a0 @p_16)
(assume t166.a1 @p_81)
(assume t166.a2 @p_101)
(assume t166.a3 @p_117)
(step t166.t0 (cl @p_217 (! (not @p_16) :named @p_80) @p_119 (! (not @p_81) :named @p_116) (! (not @p_101) :named @p_214)) :rule and_neg)
(step t166.t1 (cl (not (! (= @p_95 (! (and @p_101 (! (not (! (>= @p_98 8) :named @p_99)) :named @p_100)) :named @p_102)) :named @p_103)) (not @p_95) @p_102) :rule equiv_pos2)
(step t166.t2 (cl (! (= @p_94 @p_101) :named @p_114)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_114 3 7))
(step t166.t3 (cl (! (= @p_92 (! (not (! (>= @p_6 @p_91) :named @p_104)) :named @p_105)) :named @p_113)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_113 3 6))
(step t166.t4 (cl (! (= @p_90 (! (+ 1 @p_15) :named @p_109)) :named @p_112)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_112 3 7))
(step t166.t5 (cl (= @p_91 (! (* 8 @p_109) :named @p_110))) :rule cong :premises (t148 t166.t4))
(step t166.t6 (cl (! (= @p_110 (! (+ 8 @p_93) :named @p_106)) :named @p_111)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_111 3 7))
(step t166.t7 (cl (= @p_91 @p_106)) :rule trans :premises (t166.t5 t166.t6))
(step t166.t8 (cl (= @p_104 (! (>= @p_6 @p_106) :named @p_107))) :rule cong :premises (t96 t166.t7))
(step t166.t9 (cl (! (= @p_107 @p_99) :named @p_108)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_108 3 7))
(step t166.t10 (cl (= @p_104 @p_99)) :rule trans :premises (t166.t8 t166.t9))
(step t166.t11 (cl (= @p_105 @p_100)) :rule cong :premises (t166.t10))
(step t166.t12 (cl (= @p_92 @p_100)) :rule trans :premises (t166.t3 t166.t11))
(step t166.t13 (cl @p_103) :rule cong :premises (t166.t2 t166.t12))
(step t166.t14 (cl @p_95) :rule and :premises (t115) :args (1))
(step t166.t15 (cl @p_102) :rule resolution :premises (t166.t1 t166.t13 t166.t14))
(step t166.t16 (cl @p_101) :rule and :premises (t166.t15) :args (0))
(step t166.t17 (cl @p_217) :rule resolution :premises (t166.t0 t166.a0 t166.a3 t166.a1 t166.t16))
(step t166 (cl @p_80 @p_116 @p_214 @p_119 @p_217) :rule subproof :discharge (t166.a0 t166.a1 t166.a2 t166.a3))
(step t167 (cl (! (not @p_215) :named @p_216) @p_16) :rule and_pos :args (0))
(step t168 (cl @p_216 @p_81) :rule and_pos :args (1))
(step t169 (cl @p_216 @p_101) :rule and_pos :args (2))
(step t170 (cl @p_216 @p_117) :rule and_pos :args (3))
(step t171 (cl @p_217 @p_216 @p_216 @p_216 @p_216) :rule resolution :premises (t166 t167 t168 t169 t170))
(step t172 (cl @p_216 @p_216 @p_216 @p_216 @p_217) :rule reordering :premises (t171))
(step t173 (cl @p_216 @p_217) :rule contraction :premises (t172))
(step t174 (cl @p_251 @p_217) :rule resolution :premises (t165 t173))
(step t175 (cl @p_251 (! (not @p_217) :named @p_218)) :rule implies_neg2)
(step t176 (cl @p_251 @p_251) :rule resolution :premises (t174 t175))
(step t177 (cl @p_251) :rule contraction :premises (t176))
(step t178 (cl @p_216 @p_217) :rule implies :premises (t177))
(step t179 (cl (! (=> @p_217 false) :named @p_219) @p_217) :rule implies_neg1)
(anchor :step t180)
(assume t180.a0 @p_16)
(assume t180.a1 @p_117)
(assume t180.a2 @p_81)
(assume t180.a3 @p_101)
(step t180.t0 (cl (not @p_103) (not @p_95) @p_102) :rule equiv_pos2)
(step t180.t1 (cl @p_114) :rule hole :args ("TRUST_THEORY_REWRITE" @p_114 3 7))
(step t180.t2 (cl @p_113) :rule hole :args ("TRUST_THEORY_REWRITE" @p_113 3 6))
(step t180.t3 (cl @p_112) :rule hole :args ("TRUST_THEORY_REWRITE" @p_112 3 7))
(step t180.t4 (cl (= @p_91 @p_110)) :rule cong :premises (t148 t180.t3))
(step t180.t5 (cl @p_111) :rule hole :args ("TRUST_THEORY_REWRITE" @p_111 3 7))
(step t180.t6 (cl (= @p_91 @p_106)) :rule trans :premises (t180.t4 t180.t5))
(step t180.t7 (cl (= @p_104 @p_107)) :rule cong :premises (t96 t180.t6))
(step t180.t8 (cl @p_108) :rule hole :args ("TRUST_THEORY_REWRITE" @p_108 3 7))
(step t180.t9 (cl (= @p_104 @p_99)) :rule trans :premises (t180.t7 t180.t8))
(step t180.t10 (cl (= @p_105 @p_100)) :rule cong :premises (t180.t9))
(step t180.t11 (cl (= @p_92 @p_100)) :rule trans :premises (t180.t2 t180.t10))
(step t180.t12 (cl @p_103) :rule cong :premises (t180.t1 t180.t11))
(step t180.t13 (cl @p_95) :rule and :premises (t115) :args (1))
(step t180.t14 (cl @p_102) :rule resolution :premises (t180.t0 t180.t12 t180.t13))
(step t180.t15 (cl @p_101) :rule and :premises (t180.t14) :args (0))
(step t180.t16 (cl (not (! (= (! (< @p_98 0) :named @p_220) @p_214) :named @p_250)) (not @p_220) @p_214) :rule equiv_pos2)
(step t180.t17 (cl @p_250) :rule hole :args ("TRUST_THEORY_REWRITE" @p_250 3 6))
(step t180.t18 (cl (not (! (= @p_214 @p_220) :named @p_249)) (not @p_214) @p_220) :rule equiv_pos2)
(step t180.t19 (cl @p_249) :rule symm :premises (t180.t17))
(step t180.t20 (cl (! (=> @p_101 false) :named @p_221) @p_101) :rule implies_neg1)
(anchor :step t180.t21)
(assume t180.t21.a0 @p_101)
(step t180.t21.t0 (cl (not (! (= (! (<= (! (+ (! (* -1 @p_98) :named @p_226) (! (* 8 (! (+ @p_17 (* -1 @p_15)) :named @p_142)) :named @p_225) (! (* -8 @p_17) :named @p_202) @p_6) :named @p_227) (! (+ @p_136 (! (* 8 0) :named @p_223) (! (* -8 1) :named @p_222) 3) :named @p_224)) :named @p_228) false) :named @p_236)) (not @p_228) false) :rule equiv_pos2)
(step t180.t21.t1 (cl (! (= @p_226 (! (+ @p_140 @p_93) :named @p_244)) :named @p_248)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_248 3 7))
(step t180.t21.t2 (cl (! (= @p_225 (! (+ (! (* 8 @p_17) :named @p_141) @p_97) :named @p_243)) :named @p_247)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_247 3 7))
(step t180.t21.t3 (cl (= @p_202 @p_202)) :rule refl)
(step t180.t21.t4 (cl (= @p_227 (! (+ @p_244 @p_243 @p_202 @p_6) :named @p_245))) :rule cong :premises (t180.t21.t1 t180.t21.t2 t180.t21.t3 t96))
(step t180.t21.t5 (cl (! (= @p_245 0) :named @p_246)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_246 3 7))
(step t180.t21.t6 (cl (= @p_227 0)) :rule trans :premises (t180.t21.t4 t180.t21.t5))
(step t180.t21.t7 (cl @p_199) :rule hole :args ("TRUST_THEORY_REWRITE" @p_199 3 6))
(step t180.t21.t8 (cl (! (= @p_223 0) :named @p_242)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_242 3 6))
(step t180.t21.t9 (cl (! (= @p_222 -8) :named @p_241)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_241 3 7))
(step t180.t21.t10 (cl (= @p_224 (! (+ 0 0 -8 3) :named @p_239))) :rule cong :premises (t180.t21.t7 t180.t21.t8 t180.t21.t9 t146))
(step t180.t21.t11 (cl (! (= @p_239 -5) :named @p_240)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_240 3 7))
(step t180.t21.t12 (cl (= @p_224 -5)) :rule trans :premises (t180.t21.t10 t180.t21.t11))
(step t180.t21.t13 (cl (= @p_228 (! (<= 0 -5) :named @p_237))) :rule cong :premises (t180.t21.t6 t180.t21.t12))
(step t180.t21.t14 (cl (! (= @p_237 false) :named @p_238)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_238 3 7))
(step t180.t21.t15 (cl @p_236) :rule trans :premises (t180.t21.t13 t180.t21.t14))
(step t180.t21.t16 (cl (not (! (<= @p_226 @p_136) :named @p_234)) (not (! (= @p_225 @p_223) :named @p_232)) (not (! (<= @p_202 @p_222) :named @p_230)) @p_80 @p_228) :rule la_generic :args (1/1 -1/1 1/1 -1/1 1/1))
(step t180.t21.t17 (cl (=> (! (and @p_147 @p_101) :named @p_235) @p_234)) :rule la_mult_neg)
(step t180.t21.t18 (cl (not @p_235) @p_234) :rule implies :premises (t180.t21.t17))
(step t180.t21.t19 (cl @p_235 @p_150 @p_214) :rule and_neg)
(step t180.t21.t20 (cl (= @p_149 @p_147)) :rule equiv_simplify)
(step t180.t21.t21 (cl (not @p_149) @p_147) :rule equiv1 :premises (t180.t21.t20))
(step t180.t21.t22 (cl @p_149) :rule rare_rewrite :args ("evaluate"))
(step t180.t21.t23 (cl @p_147) :rule resolution :premises (t180.t21.t21 t180.t21.t22))
(step t180.t21.t24 (cl @p_235) :rule resolution :premises (t180.t21.t19 t180.t21.t23 t180.t15))
(step t180.t21.t25 (cl @p_234) :rule resolution :premises (t180.t21.t18 t180.t21.t24))
(step t180.t21.t26 (cl (=> (! (and (! (> 8 0) :named @p_153) (! (= @p_142 0) :named @p_176)) :named @p_233) @p_232)) :rule la_mult_pos)
(step t180.t21.t27 (cl (not @p_233) @p_232) :rule implies :premises (t180.t21.t26))
(step t180.t21.t28 (cl @p_233 (! (not @p_153) :named @p_173) (! (not @p_176) :named @p_185)) :rule and_neg)
(step t180.t21.t29 (cl (= (! (= @p_153 true) :named @p_172) @p_153)) :rule equiv_simplify)
(step t180.t21.t30 (cl (not @p_172) @p_153) :rule equiv1 :premises (t180.t21.t29))
(step t180.t21.t31 (cl @p_172) :rule rare_rewrite :args ("evaluate"))
(step t180.t21.t32 (cl @p_153) :rule resolution :premises (t180.t21.t30 t180.t21.t31))
(step t180.t21.t33 (cl (! (not (! (= @p_81 @p_176) :named @p_180)) :named @p_182) @p_116 @p_176) :rule equiv_pos2)
(step t180.t21.t34 (cl (! (= @p_176 @p_81) :named @p_181)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_181 3 7))
(step t180.t21.t35 (cl @p_180) :rule symm :premises (t180.t21.t34))
(step t180.t21.t36 (cl @p_176) :rule resolution :premises (t180.t21.t33 t180.t21.t35 t180.a2))
(step t180.t21.t37 (cl @p_233) :rule resolution :premises (t180.t21.t28 t180.t21.t32 t180.t21.t36))
(step t180.t21.t38 (cl @p_232) :rule resolution :premises (t180.t21.t27 t180.t21.t37))
(step t180.t21.t39 (cl (=> (! (and (! (< -8 0) :named @p_177) @p_117) :named @p_231) @p_230)) :rule la_mult_neg)
(step t180.t21.t40 (cl (not @p_231) @p_230) :rule implies :premises (t180.t21.t39))
(step t180.t21.t41 (cl @p_231 (! (not @p_177) :named @p_184) @p_119) :rule and_neg)
(step t180.t21.t42 (cl (= (! (= @p_177 true) :named @p_183) @p_177)) :rule equiv_simplify)
(step t180.t21.t43 (cl (not @p_183) @p_177) :rule equiv1 :premises (t180.t21.t42))
(step t180.t21.t44 (cl @p_183) :rule rare_rewrite :args ("evaluate"))
(step t180.t21.t45 (cl @p_177) :rule resolution :premises (t180.t21.t43 t180.t21.t44))
(step t180.t21.t46 (cl @p_231) :rule resolution :premises (t180.t21.t41 t180.t21.t45 t180.a1))
(step t180.t21.t47 (cl @p_230) :rule resolution :premises (t180.t21.t40 t180.t21.t46))
(step t180.t21.t48 (cl @p_228) :rule resolution :premises (t180.t21.t16 t180.t21.t25 t180.t21.t38 t180.t21.t47 t180.a0))
(step t180.t21.t49 (cl false) :rule resolution :premises (t180.t21.t0 t180.t21.t15 t180.t21.t48))
(step t180.t21 (cl @p_214 false) :rule subproof :discharge (t180.t21.a0))
(step t180.t22 (cl @p_221 false) :rule resolution :premises (t180.t20 t180.t21))
(step t180.t23 (cl @p_221 @p_132) :rule implies_neg2)
(step t180.t24 (cl @p_221 @p_221) :rule resolution :premises (t180.t22 t180.t23))
(step t180.t25 (cl @p_221) :rule contraction :premises (t180.t24))
(step t180.t26 (cl (= @p_221 @p_214)) :rule implies_simplify)
(step t180.t27 (cl (not @p_221) @p_214) :rule equiv1 :premises (t180.t26))
(step t180.t28 (cl @p_214) :rule resolution :premises (t180.t25 t180.t27))
(step t180.t29 (cl @p_220) :rule resolution :premises (t180.t18 t180.t19 t180.t28))
(step t180.t30 (cl @p_214) :rule resolution :premises (t180.t16 t180.t17 t180.t29))
(step t180.t31 (cl) :rule resolution :premises (t180.t15 t180.t30))
(step t180 (cl @p_80 @p_119 @p_116 @p_214 false) :rule subproof :discharge (t180.a0 t180.a1 t180.a2 t180.a3))
(step t181 (cl @p_218 @p_16) :rule and_pos :args (0))
(step t182 (cl @p_218 @p_117) :rule and_pos :args (1))
(step t183 (cl @p_218 @p_81) :rule and_pos :args (2))
(step t184 (cl @p_218 @p_101) :rule and_pos :args (3))
(step t185 (cl false @p_218 @p_218 @p_218 @p_218) :rule resolution :premises (t180 t181 t182 t183 t184))
(step t186 (cl @p_218 @p_218 @p_218 @p_218 false) :rule reordering :premises (t185))
(step t187 (cl @p_218 false) :rule contraction :premises (t186))
(step t188 (cl @p_219 false) :rule resolution :premises (t179 t187))
(step t189 (cl @p_219 @p_132) :rule implies_neg2)
(step t190 (cl @p_219 @p_219) :rule resolution :premises (t188 t189))
(step t191 (cl @p_219) :rule contraction :premises (t190))
(step t192 (cl (= @p_219 @p_218)) :rule implies_simplify)
(step t193 (cl (not @p_219) @p_218) :rule equiv1 :premises (t192))
(step t194 (cl @p_218) :rule resolution :premises (t191 t193))
(step t195 (cl @p_216) :rule resolution :premises (t178 t194))
(step t196 (cl @p_80 @p_116 @p_214 @p_119) :rule not_and :premises (t195))
(step t197 (cl (not @p_103) (not @p_95) @p_102) :rule equiv_pos2)
(step t198 (cl @p_114) :rule hole :args ("TRUST_THEORY_REWRITE" @p_114 3 7))
(step t199 (cl @p_113) :rule hole :args ("TRUST_THEORY_REWRITE" @p_113 3 6))
(step t200 (cl @p_112) :rule hole :args ("TRUST_THEORY_REWRITE" @p_112 3 7))
(step t201 (cl (= @p_91 @p_110)) :rule cong :premises (t148 t200))
(step t202 (cl @p_111) :rule hole :args ("TRUST_THEORY_REWRITE" @p_111 3 7))
(step t203 (cl (= @p_91 @p_106)) :rule trans :premises (t201 t202))
(step t204 (cl (= @p_104 @p_107)) :rule cong :premises (t96 t203))
(step t205 (cl @p_108) :rule hole :args ("TRUST_THEORY_REWRITE" @p_108 3 7))
(step t206 (cl (= @p_104 @p_99)) :rule trans :premises (t204 t205))
(step t207 (cl (= @p_105 @p_100)) :rule cong :premises (t206))
(step t208 (cl (= @p_92 @p_100)) :rule trans :premises (t199 t207))
(step t209 (cl @p_103) :rule cong :premises (t198 t208))
(step t210 (cl @p_95) :rule and :premises (t115) :args (1))
(step t211 (cl @p_102) :rule resolution :premises (t197 t209 t210))
(step t212 (cl @p_101) :rule and :premises (t211) :args (0))
(step t213 (cl (not (! (= (! (or (! (not (! (>= @p_6 0) :named @p_96)) :named @p_115) @p_122 @p_116 (! (not @p_100) :named @p_121) (! (not @p_119) :named @p_120)) :named @p_123) (! (or @p_115 @p_18 @p_116 @p_99 @p_117) :named @p_118)) :named @p_211)) (not @p_123) @p_118) :rule equiv_pos2)
(step t214 (cl (= @p_115 @p_115)) :rule refl)
(step t215 (cl (= @p_116 @p_116)) :rule refl)
(step t216 (cl (! (= @p_121 @p_99) :named @p_208)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_208 1 6))
(step t217 (cl (! (= @p_120 @p_117) :named @p_212)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_212 1 6))
(step t218 (cl @p_211) :rule cong :premises (t214 t1 t215 t216 t217))
(step t219 (cl (! (=> (! (and @p_96 @p_19 @p_81 @p_100 @p_119) :named @p_127) (! (and @p_96 @p_19 @p_119 @p_81 @p_100) :named @p_129)) :named @p_210) @p_127) :rule implies_neg1)
(anchor :step t220)
(assume t220.a0 @p_96)
(assume t220.a1 @p_19)
(assume t220.a2 @p_81)
(assume t220.a3 @p_100)
(assume t220.a4 @p_119)
(step t220.t0 (cl @p_129 @p_115 @p_122 @p_120 @p_116 @p_121) :rule and_neg)
(step t220.t1 (cl (and @p_96 (not (>= @p_6 16)))) :rule bv_intblast_bounds)
(step t220.t2 (cl @p_96) :rule and :premises (t220.t1) :args (0))
(step t220.t3 (cl (not @p_20) (not @p_5) @p_19) :rule equiv_pos2)
(step t220.t4 (cl @p_54) :rule hole :args ("TRUST_THEORY_REWRITE" @p_54 4 6))
(step t220.t5 (cl (= #b0000 #b0000)) :rule refl)
(step t220.t6 (cl (= s s)) :rule refl)
(step t220.t7 (cl @p_53) :rule hole :args ("TRUST_THEORY_REWRITE" @p_53 4 6))
(step t220.t8 (cl @p_52) :rule hole :args ("TRUST_THEORY_REWRITE" @p_52 4 6))
(step t220.t9 (cl (= @p_49 @p_50)) :rule cong :premises (t220.t6 t220.t8))
(step t220.t10 (cl @p_51) :rule hole :args ("TRUST_THEORY_REWRITE" @p_51 4 7))
(step t220.t11 (cl (= @p_49 s)) :rule trans :premises (t220.t9 t220.t10))
(step t220.t12 (cl (= @p_2 s)) :rule trans :premises (t220.t7 t220.t11))
(step t220.t13 (cl (= @p_3 @p_45)) :rule cong :premises (t220.t6 t220.t12))
(step t220.t14 (cl (= @p_48 @p_46)) :rule cong :premises (t220.t5 t220.t13))
(step t220.t15 (cl (= @p_4 @p_46)) :rule trans :premises (t220.t4 t220.t14))
(step t220.t16 (cl (= @p_5 @p_47)) :rule cong :premises (t220.t15))
(step t220.t17 (cl (= @p_47 @p_41)) :rule bv_intblast)
(step t220.t18 (cl @p_32) :rule refl)
(step t220.t19 (cl @p_31) :rule refl)
(step t220.t20 (cl @p_30) :rule refl)
(step t220.t21 (cl @p_29) :rule refl)
(step t220.t22 (cl @p_28) :rule refl)
(step t220.t23 (cl @p_44) :rule hole :args ("TRUST_THEORY_REWRITE" @p_44 3 7))
(step t220.t24 (cl (= @p_36 @p_8)) :rule cong :premises (t220.t22 t220.t23 t129))
(step t220.t25 (cl (= @p_37 @p_11)) :rule cong :premises (t220.t20 t220.t21 t220.t24))
(step t220.t26 (cl (= @p_38 @p_14)) :rule cong :premises (t220.t18 t220.t19 t220.t25))
(step t220.t27 (cl (= @p_39 @p_17)) :rule cong :premises (t113 t156 t220.t26))
(step t220.t28 (cl (= @p_40 @p_42)) :rule cong :premises (t129 t220.t27))
(step t220.t29 (cl @p_43) :rule hole :args ("TRUST_THEORY_REWRITE" @p_43 3 7))
(step t220.t30 (cl (= @p_40 @p_18)) :rule trans :premises (t220.t28 t220.t29))
(step t220.t31 (cl (= @p_41 @p_19)) :rule cong :premises (t220.t30))
(step t220.t32 (cl (= @p_18 @p_18)) :rule cong :premises (t155 t129))
(step t220.t33 (cl (= @p_19 @p_19)) :rule cong :premises (t220.t32))
(step t220.t34 (cl @p_20) :rule trans :premises (t220.t16 t220.t17 t220.t31 t220.t33))
(step t220.t35 (cl @p_19) :rule resolution :premises (t220.t3 t220.t34 a0))
(step t220.t36 (cl @p_100) :rule and :premises (t211) :args (1))
(step t220.t37 (cl @p_129) :rule resolution :premises (t220.t0 t220.t2 t220.t35 t220.a4 t220.a2 t220.t36))
(step t220 (cl @p_115 @p_122 @p_116 @p_121 @p_120 @p_129) :rule subproof :discharge (t220.a0 t220.a1 t220.a2 t220.a3 t220.a4))
(step t221 (cl (! (not @p_127) :named @p_128) @p_96) :rule and_pos :args (0))
(step t222 (cl @p_128 @p_19) :rule and_pos :args (1))
(step t223 (cl @p_128 @p_81) :rule and_pos :args (2))
(step t224 (cl @p_128 @p_100) :rule and_pos :args (3))
(step t225 (cl @p_128 @p_119) :rule and_pos :args (4))
(step t226 (cl @p_129 @p_128 @p_128 @p_128 @p_128 @p_128) :rule resolution :premises (t220 t221 t222 t223 t224 t225))
(step t227 (cl @p_128 @p_128 @p_128 @p_128 @p_128 @p_129) :rule reordering :premises (t226))
(step t228 (cl @p_128 @p_129) :rule contraction :premises (t227))
(step t229 (cl @p_210 @p_129) :rule resolution :premises (t219 t228))
(step t230 (cl @p_210 (! (not @p_129) :named @p_130)) :rule implies_neg2)
(step t231 (cl @p_210 @p_210) :rule resolution :premises (t229 t230))
(step t232 (cl @p_210) :rule contraction :premises (t231))
(step t233 (cl @p_128 @p_129) :rule implies :premises (t232))
(step t234 (cl (! (=> @p_129 false) :named @p_131) @p_129) :rule implies_neg1)
(anchor :step t235)
(assume t235.a0 @p_96)
(assume t235.a1 @p_19)
(assume t235.a2 @p_119)
(assume t235.a3 @p_81)
(assume t235.a4 @p_100)
(step t235.t0 (cl (! (not (! (= @p_100 (! (< @p_98 8) :named @p_133)) :named @p_186)) :named @p_188) @p_121 @p_133) :rule equiv_pos2)
(step t235.t1 (cl (! (= @p_133 @p_100) :named @p_187)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_187 3 6))
(step t235.t2 (cl @p_186) :rule symm :premises (t235.t1))
(step t235.t3 (cl @p_100) :rule and :premises (t211) :args (1))
(step t235.t4 (cl @p_133) :rule resolution :premises (t235.t0 t235.t2 t235.t3))
(step t235.t5 (cl (not (! (= @p_99 (! (not @p_133) :named @p_134)) :named @p_209)) @p_100 @p_134) :rule equiv_pos2)
(step t235.t6 (cl (= @p_134 @p_121)) :rule cong :premises (t235.t1))
(step t235.t7 (cl (! (= @p_134 @p_99) :named @p_207)) :rule trans :premises (t235.t6 t216))
(step t235.t8 (cl @p_209) :rule symm :premises (t235.t7))
(step t235.t9 (cl (not @p_207) (not @p_134) @p_99) :rule equiv_pos2)
(step t235.t10 (cl (! (=> @p_133 false) :named @p_135) @p_133) :rule implies_neg1)
(anchor :step t235.t11)
(assume t235.t11.a0 @p_133)
(step t235.t11.t0 (cl (not (! (= (! (< (! (+ @p_98 (! (* -8 @p_142) :named @p_143) @p_141 @p_140) :named @p_144) (! (+ 8 (! (* -8 0) :named @p_138) (! (* 8 -1) :named @p_137) @p_136) :named @p_139)) :named @p_145) false) :named @p_190)) (not @p_145) false) :rule equiv_pos2)
(step t235.t11.t1 (cl (! (= @p_145 (! (not (! (>= @p_144 @p_139) :named @p_191)) :named @p_192)) :named @p_206)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_206 3 6))
(step t235.t11.t2 (cl (! (= @p_144 (! (+ @p_141 @p_140 @p_93 @p_97 @p_202 @p_6) :named @p_203)) :named @p_205)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_205 3 6))
(step t235.t11.t3 (cl (! (= @p_203 0) :named @p_204)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_204 3 7))
(step t235.t11.t4 (cl (= @p_144 0)) :rule trans :premises (t235.t11.t2 t235.t11.t3))
(step t235.t11.t5 (cl (! (= @p_138 0) :named @p_201)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_201 3 6))
(step t235.t11.t6 (cl (! (= @p_137 -8) :named @p_200)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_200 3 7))
(step t235.t11.t7 (cl @p_199) :rule hole :args ("TRUST_THEORY_REWRITE" @p_199 3 6))
(step t235.t11.t8 (cl (= @p_139 (! (+ 8 0 -8 0) :named @p_197))) :rule cong :premises (t148 t235.t11.t5 t235.t11.t6 t235.t11.t7))
(step t235.t11.t9 (cl (! (= @p_197 0) :named @p_198)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_198 3 7))
(step t235.t11.t10 (cl (= @p_139 0)) :rule trans :premises (t235.t11.t8 t235.t11.t9))
(step t235.t11.t11 (cl (= @p_191 @p_195)) :rule cong :premises (t235.t11.t4 t235.t11.t10))
(step t235.t11.t12 (cl @p_196) :rule hole :args ("TRUST_THEORY_REWRITE" @p_196 3 7))
(step t235.t11.t13 (cl (= @p_191 true)) :rule trans :premises (t235.t11.t11 t235.t11.t12))
(step t235.t11.t14 (cl (= @p_192 @p_193)) :rule cong :premises (t235.t11.t13))
(step t235.t11.t15 (cl @p_194) :rule hole :args ("TRUST_THEORY_REWRITE" @p_194 1 7))
(step t235.t11.t16 (cl (= @p_192 false)) :rule trans :premises (t235.t11.t14 t235.t11.t15))
(step t235.t11.t17 (cl @p_190) :rule trans :premises (t235.t11.t1 t235.t11.t16))
(step t235.t11.t18 (cl @p_134 (not (! (= @p_143 @p_138) :named @p_175)) (not (! (<= @p_141 @p_137) :named @p_151)) (! (not (! (<= @p_140 @p_136) :named @p_146)) :named @p_189) @p_145) :rule la_generic :args (1/1 -1/1 1/1 1/1 1/1))
(step t235.t11.t19 (cl @p_188 @p_121 @p_133) :rule equiv_pos2)
(step t235.t11.t20 (cl @p_100) :rule and :premises (t211) :args (1))
(step t235.t11.t21 (cl @p_133) :rule resolution :premises (t235.t11.t19 t235.t2 t235.t11.t20))
(step t235.t11.t22 (cl (=> (! (and @p_177 @p_176) :named @p_178) @p_175)) :rule la_mult_neg)
(step t235.t11.t23 (cl (not @p_178) @p_175) :rule implies :premises (t235.t11.t22))
(step t235.t11.t24 (cl @p_178 @p_184 @p_185) :rule and_neg)
(step t235.t11.t25 (cl (= @p_183 @p_177)) :rule equiv_simplify)
(step t235.t11.t26 (cl (not @p_183) @p_177) :rule equiv1 :premises (t235.t11.t25))
(step t235.t11.t27 (cl @p_183) :rule rare_rewrite :args ("evaluate"))
(step t235.t11.t28 (cl @p_177) :rule resolution :premises (t235.t11.t26 t235.t11.t27))
(step t235.t11.t29 (cl @p_182 @p_116 @p_176) :rule equiv_pos2)
(step t235.t11.t30 (cl @p_181) :rule hole :args ("TRUST_THEORY_REWRITE" @p_181 3 7))
(step t235.t11.t31 (cl @p_180) :rule symm :premises (t235.t11.t30))
(step t235.t11.t32 (cl @p_176) :rule resolution :premises (t235.t11.t29 t235.t11.t31 t235.a3))
(step t235.t11.t33 (cl @p_178) :rule resolution :premises (t235.t11.t24 t235.t11.t28 t235.t11.t32))
(step t235.t11.t34 (cl @p_175) :rule resolution :premises (t235.t11.t23 t235.t11.t33))
(step t235.t11.t35 (cl (=> (! (and @p_153 (! (<= @p_17 -1) :named @p_152)) :named @p_154) @p_151)) :rule la_mult_pos)
(step t235.t11.t36 (cl (not @p_154) @p_151) :rule implies :premises (t235.t11.t35))
(step t235.t11.t37 (cl @p_154 @p_173 (! (not @p_152) :named @p_174)) :rule and_neg)
(step t235.t11.t38 (cl (= @p_172 @p_153)) :rule equiv_simplify)
(step t235.t11.t39 (cl (not @p_172) @p_153) :rule equiv1 :premises (t235.t11.t38))
(step t235.t11.t40 (cl @p_172) :rule rare_rewrite :args ("evaluate"))
(step t235.t11.t41 (cl @p_153) :rule resolution :premises (t235.t11.t39 t235.t11.t40))
(step t235.t11.t42 (cl @p_171 @p_152) :rule la_generic :args (1/1 1/1))
(step t235.t11.t43 (cl @p_170) :rule la_disequality)
(step t235.t11.t44 (cl @p_18 @p_169 @p_157) :rule or :premises (t235.t11.t43))
(step t235.t11.t45 (cl (not @p_20) (not @p_5) @p_19) :rule equiv_pos2)
(step t235.t11.t46 (cl @p_54) :rule hole :args ("TRUST_THEORY_REWRITE" @p_54 4 6))
(step t235.t11.t47 (cl (= #b0000 #b0000)) :rule refl)
(step t235.t11.t48 (cl (= s s)) :rule refl)
(step t235.t11.t49 (cl @p_53) :rule hole :args ("TRUST_THEORY_REWRITE" @p_53 4 6))
(step t235.t11.t50 (cl @p_52) :rule hole :args ("TRUST_THEORY_REWRITE" @p_52 4 6))
(step t235.t11.t51 (cl (= @p_49 @p_50)) :rule cong :premises (t235.t11.t48 t235.t11.t50))
(step t235.t11.t52 (cl @p_51) :rule hole :args ("TRUST_THEORY_REWRITE" @p_51 4 7))
(step t235.t11.t53 (cl (= @p_49 s)) :rule trans :premises (t235.t11.t51 t235.t11.t52))
(step t235.t11.t54 (cl (= @p_2 s)) :rule trans :premises (t235.t11.t49 t235.t11.t53))
(step t235.t11.t55 (cl (= @p_3 @p_45)) :rule cong :premises (t235.t11.t48 t235.t11.t54))
(step t235.t11.t56 (cl (= @p_48 @p_46)) :rule cong :premises (t235.t11.t47 t235.t11.t55))
(step t235.t11.t57 (cl (= @p_4 @p_46)) :rule trans :premises (t235.t11.t46 t235.t11.t56))
(step t235.t11.t58 (cl (= @p_5 @p_47)) :rule cong :premises (t235.t11.t57))
(step t235.t11.t59 (cl (= @p_47 @p_41)) :rule bv_intblast)
(step t235.t11.t60 (cl @p_32) :rule refl)
(step t235.t11.t61 (cl @p_31) :rule refl)
(step t235.t11.t62 (cl @p_30) :rule refl)
(step t235.t11.t63 (cl @p_29) :rule refl)
(step t235.t11.t64 (cl @p_28) :rule refl)
(step t235.t11.t65 (cl @p_44) :rule hole :args ("TRUST_THEORY_REWRITE" @p_44 3 7))
(step t235.t11.t66 (cl (= @p_36 @p_8)) :rule cong :premises (t235.t11.t64 t235.t11.t65 t129))
(step t235.t11.t67 (cl (= @p_37 @p_11)) :rule cong :premises (t235.t11.t62 t235.t11.t63 t235.t11.t66))
(step t235.t11.t68 (cl (= @p_38 @p_14)) :rule cong :premises (t235.t11.t60 t235.t11.t61 t235.t11.t67))
(step t235.t11.t69 (cl (= @p_39 @p_17)) :rule cong :premises (t113 t156 t235.t11.t68))
(step t235.t11.t70 (cl (= @p_40 @p_42)) :rule cong :premises (t129 t235.t11.t69))
(step t235.t11.t71 (cl @p_43) :rule hole :args ("TRUST_THEORY_REWRITE" @p_43 3 7))
(step t235.t11.t72 (cl (= @p_40 @p_18)) :rule trans :premises (t235.t11.t70 t235.t11.t71))
(step t235.t11.t73 (cl (= @p_41 @p_19)) :rule cong :premises (t235.t11.t72))
(step t235.t11.t74 (cl (= @p_18 @p_18)) :rule cong :premises (t155 t129))
(step t235.t11.t75 (cl (= @p_19 @p_19)) :rule cong :premises (t235.t11.t74))
(step t235.t11.t76 (cl @p_20) :rule trans :premises (t235.t11.t58 t235.t11.t59 t235.t11.t73 t235.t11.t75))
(step t235.t11.t77 (cl @p_19) :rule resolution :premises (t235.t11.t45 t235.t11.t76 a0))
(step t235.t11.t78 (cl (! (not (! (< @p_17 1) :named @p_162)) :named @p_168) @p_161) :rule la_generic :args (1/1 1/1))
(step t235.t11.t79 (cl (! (not (! (= @p_119 @p_162) :named @p_165)) :named @p_167) @p_120 @p_162) :rule equiv_pos2)
(step t235.t11.t80 (cl (! (= @p_162 @p_119) :named @p_166)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_166 3 6))
(step t235.t11.t81 (cl @p_165) :rule symm :premises (t235.t11.t80))
(step t235.t11.t82 (cl @p_162) :rule resolution :premises (t235.t11.t79 t235.t11.t81 t235.a2))
(step t235.t11.t83 (cl @p_161) :rule resolution :premises (t235.t11.t78 t235.t11.t82))
(step t235.t11.t84 (cl (! (not (! (= @p_155 @p_157) :named @p_158)) :named @p_159) @p_155 (! (not @p_157) :named @p_160)) :rule equiv_pos1)
(step t235.t11.t85 (cl @p_158) :rule comp_simplify)
(step t235.t11.t86 (cl @p_155) :rule resolution :premises (t235.t11.t44 t235.t11.t77 t235.t11.t83 t235.t11.t84 t235.t11.t85))
(step t235.t11.t87 (cl @p_152) :rule resolution :premises (t235.t11.t42 t235.t11.t86))
(step t235.t11.t88 (cl @p_154) :rule resolution :premises (t235.t11.t37 t235.t11.t41 t235.t11.t87))
(step t235.t11.t89 (cl @p_151) :rule resolution :premises (t235.t11.t36 t235.t11.t88))
(step t235.t11.t90 (cl (=> (! (and @p_147 @p_96) :named @p_148) @p_146)) :rule la_mult_neg)
(step t235.t11.t91 (cl (not @p_148) @p_146) :rule implies :premises (t235.t11.t90))
(step t235.t11.t92 (cl @p_148 @p_150 @p_115) :rule and_neg)
(step t235.t11.t93 (cl (= @p_149 @p_147)) :rule equiv_simplify)
(step t235.t11.t94 (cl (not @p_149) @p_147) :rule equiv1 :premises (t235.t11.t93))
(step t235.t11.t95 (cl @p_149) :rule rare_rewrite :args ("evaluate"))
(step t235.t11.t96 (cl @p_147) :rule resolution :premises (t235.t11.t94 t235.t11.t95))
(step t235.t11.t97 (cl (and @p_96 (not (>= @p_6 16)))) :rule bv_intblast_bounds)
(step t235.t11.t98 (cl @p_96) :rule and :premises (t235.t11.t97) :args (0))
(step t235.t11.t99 (cl @p_148) :rule resolution :premises (t235.t11.t92 t235.t11.t96 t235.t11.t98))
(step t235.t11.t100 (cl @p_146) :rule resolution :premises (t235.t11.t91 t235.t11.t99))
(step t235.t11.t101 (cl @p_145) :rule resolution :premises (t235.t11.t18 t235.t11.t21 t235.t11.t34 t235.t11.t89 t235.t11.t100))
(step t235.t11.t102 (cl false) :rule resolution :premises (t235.t11.t0 t235.t11.t17 t235.t11.t101))
(step t235.t11 (cl @p_134 false) :rule subproof :discharge (t235.t11.a0))
(step t235.t12 (cl @p_135 false) :rule resolution :premises (t235.t10 t235.t11))
(step t235.t13 (cl @p_135 @p_132) :rule implies_neg2)
(step t235.t14 (cl @p_135 @p_135) :rule resolution :premises (t235.t12 t235.t13))
(step t235.t15 (cl @p_135) :rule contraction :premises (t235.t14))
(step t235.t16 (cl (= @p_135 @p_134)) :rule implies_simplify)
(step t235.t17 (cl (not @p_135) @p_134) :rule equiv1 :premises (t235.t16))
(step t235.t18 (cl @p_134) :rule resolution :premises (t235.t15 t235.t17))
(step t235.t19 (cl @p_99) :rule resolution :premises (t235.t9 t235.t7 t235.t18))
(step t235.t20 (cl @p_134) :rule resolution :premises (t235.t5 t235.t8 t235.t19))
(step t235.t21 (cl) :rule resolution :premises (t235.t4 t235.t20))
(step t235 (cl @p_115 @p_122 @p_120 @p_116 @p_121 false) :rule subproof :discharge (t235.a0 t235.a1 t235.a2 t235.a3 t235.a4))
(step t236 (cl @p_130 @p_96) :rule and_pos :args (0))
(step t237 (cl @p_130 @p_19) :rule and_pos :args (1))
(step t238 (cl @p_130 @p_119) :rule and_pos :args (2))
(step t239 (cl @p_130 @p_81) :rule and_pos :args (3))
(step t240 (cl @p_130 @p_100) :rule and_pos :args (4))
(step t241 (cl false @p_130 @p_130 @p_130 @p_130 @p_130) :rule resolution :premises (t235 t236 t237 t238 t239 t240))
(step t242 (cl @p_130 @p_130 @p_130 @p_130 @p_130 false) :rule reordering :premises (t241))
(step t243 (cl @p_130 false) :rule contraction :premises (t242))
(step t244 (cl @p_131 false) :rule resolution :premises (t234 t243))
(step t245 (cl @p_131 @p_132) :rule implies_neg2)
(step t246 (cl @p_131 @p_131) :rule resolution :premises (t244 t245))
(step t247 (cl @p_131) :rule contraction :premises (t246))
(step t248 (cl (= @p_131 @p_130)) :rule implies_simplify)
(step t249 (cl (not @p_131) @p_130) :rule equiv1 :premises (t248))
(step t250 (cl @p_130) :rule resolution :premises (t247 t249))
(step t251 (cl @p_128) :rule resolution :premises (t233 t250))
(step t252 (cl @p_115 @p_122 @p_116 @p_121 @p_120) :rule not_and :premises (t251))
(step t253 (cl @p_123 (! (not @p_115) :named @p_126)) :rule or_neg :args (0))
(step t254 (cl @p_123 @p_125) :rule or_neg :args (1))
(step t255 (cl @p_123 (not @p_116)) :rule or_neg :args (2))
(step t256 (cl @p_123 (not @p_121)) :rule or_neg :args (3))
(step t257 (cl @p_123 (! (not @p_120) :named @p_124)) :rule or_neg :args (4))
(step t258 (cl @p_123 @p_123 @p_123 @p_123 @p_123) :rule resolution :premises (t252 t253 t254 t255 t256 t257))
(step t259 (cl @p_123) :rule contraction :premises (t258))
(step t260 (cl @p_118) :rule resolution :premises (t213 t218 t259))
(step t261 (cl @p_115 @p_18 @p_116 @p_99 @p_117) :rule or :premises (t260))
(step t262 (cl @p_18 @p_99 @p_115 @p_116 @p_117) :rule reordering :premises (t261))
(step t263 (cl @p_100) :rule and :premises (t211) :args (1))
(step t264 (cl (and @p_96 (not (>= @p_6 16)))) :rule bv_intblast_bounds)
(step t265 (cl @p_96) :rule and :premises (t264) :args (0))
(step t266 (cl (not @p_20) (not @p_5) @p_19) :rule equiv_pos2)
(step t267 (cl @p_54) :rule hole :args ("TRUST_THEORY_REWRITE" @p_54 4 6))
(step t268 (cl (= #b0000 #b0000)) :rule refl)
(step t269 (cl (= s s)) :rule refl)
(step t270 (cl @p_53) :rule hole :args ("TRUST_THEORY_REWRITE" @p_53 4 6))
(step t271 (cl @p_52) :rule hole :args ("TRUST_THEORY_REWRITE" @p_52 4 6))
(step t272 (cl (= @p_49 @p_50)) :rule cong :premises (t269 t271))
(step t273 (cl @p_51) :rule hole :args ("TRUST_THEORY_REWRITE" @p_51 4 7))
(step t274 (cl (= @p_49 s)) :rule trans :premises (t272 t273))
(step t275 (cl (= @p_2 s)) :rule trans :premises (t270 t274))
(step t276 (cl (= @p_3 @p_45)) :rule cong :premises (t269 t275))
(step t277 (cl (= @p_48 @p_46)) :rule cong :premises (t268 t276))
(step t278 (cl (= @p_4 @p_46)) :rule trans :premises (t267 t277))
(step t279 (cl (= @p_5 @p_47)) :rule cong :premises (t278))
(step t280 (cl (= @p_47 @p_41)) :rule bv_intblast)
(step t281 (cl @p_32) :rule refl)
(step t282 (cl @p_31) :rule refl)
(step t283 (cl @p_30) :rule refl)
(step t284 (cl @p_29) :rule refl)
(step t285 (cl @p_28) :rule refl)
(step t286 (cl @p_44) :rule hole :args ("TRUST_THEORY_REWRITE" @p_44 3 7))
(step t287 (cl (= @p_36 @p_8)) :rule cong :premises (t285 t286 t129))
(step t288 (cl (= @p_37 @p_11)) :rule cong :premises (t283 t284 t287))
(step t289 (cl (= @p_38 @p_14)) :rule cong :premises (t281 t282 t288))
(step t290 (cl (= @p_39 @p_17)) :rule cong :premises (t113 t156 t289))
(step t291 (cl (= @p_40 @p_42)) :rule cong :premises (t129 t290))
(step t292 (cl @p_43) :rule hole :args ("TRUST_THEORY_REWRITE" @p_43 3 7))
(step t293 (cl (= @p_40 @p_18)) :rule trans :premises (t291 t292))
(step t294 (cl (= @p_41 @p_19)) :rule cong :premises (t293))
(step t295 (cl (= @p_18 @p_18)) :rule cong :premises (t155 t129))
(step t296 (cl (= @p_19 @p_19)) :rule cong :premises (t295))
(step t297 (cl @p_20) :rule trans :premises (t279 t280 t294 t296))
(step t298 (cl @p_19) :rule resolution :premises (t266 t297 a0))
(step t299 (cl @p_80 @p_81) :rule ite2 :premises (t163))
(step t300 (cl @p_81 @p_80) :rule reordering :premises (t299))
(step t301 (cl @p_80) :rule resolution :premises (t196 t212 t262 t263 t265 t298 t300))
(step t302 (cl @p_79) :rule resolution :premises (t164 t301))
(step t303 (cl (! (not (! (= (! (ite @p_13 (! (= @p_14 @p_12) :named @p_254) @p_252) :named @p_255) @p_255) :named @p_256)) :named @p_259) (! (not @p_255) :named @p_260) @p_255) :rule equiv_pos2)
(step t304 (cl @p_24) :rule refl)
(step t305 (cl (and @p_31 (! (and (! (<= (! (* 4 @p_12) :named @p_266) @p_6) :named @p_267) (! (< @p_6 (! (* 4 (! (+ @p_12 1) :named @p_263)) :named @p_264)) :named @p_265)) :named @p_268))) :rule hole :args ("ARITH_REDUCTION" @p_12))
(step t306 (cl @p_31) :rule and :premises (t305) :args (0))
(step t307 (cl (! (= @p_254 @p_254) :named @p_258)) :rule cong :premises (t304 t306))
(step t308 (cl (= (! (= @p_25 true) :named @p_261) @p_25)) :rule equiv_simplify)
(step t309 (cl (not @p_261) @p_25) :rule equiv1 :premises (t308))
(step t310 (cl (! (= @p_25 @p_25) :named @p_262)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_262 3 7))
(step t311 (cl @p_25) :rule refl)
(step t312 (cl @p_25) :rule trans :premises (t311 t132))
(step t313 (cl @p_262) :rule cong :premises (t312 t132))
(step t314 (cl @p_261) :rule hole :args ("TRUST_THEORY_REWRITE" @p_261 3 6))
(step t315 (cl @p_261) :rule trans :premises (t310 t313 t314))
(step t316 (cl @p_25) :rule resolution :premises (t309 t315))
(step t317 (cl (! (= @p_252 @p_252) :named @p_257)) :rule cong :premises (t304 t316))
(step t318 (cl @p_256) :rule cong :premises (t281 t307 t317))
(step t319 (cl @p_259 @p_260 @p_255) :rule equiv_pos2)
(step t320 (cl @p_258) :rule cong :premises (t138 t282))
(step t321 (cl @p_25) :rule refl)
(step t322 (cl @p_257) :rule cong :premises (t138 t321))
(step t323 (cl @p_256) :rule cong :premises (t281 t320 t322))
(step t324 (cl @p_255) :rule hole)
(step t325 (cl @p_255) :rule resolution :premises (t319 t323 t324))
(step t326 (cl @p_255) :rule resolution :premises (t303 t318 t325))
(step t327 (cl @p_13 @p_252) :rule ite1 :premises (t326))
(step t328 (cl (not (! (= (! (or @p_312 (! (not @p_254) :named @p_313) (! (not (! (>= (! (+ @p_6 (! (* -4 @p_12) :named @p_294)) :named @p_295) 0) :named @p_298)) :named @p_384) @p_119 (! (not (! (not (! (>= @p_6 4) :named @p_269)) :named @p_270)) :named @p_272)) :named @p_386) (! (or @p_312 @p_313 @p_384 @p_119 @p_269) :named @p_385)) :named @p_431)) (not @p_386) @p_385) :rule equiv_pos2)
(step t329 (cl (= @p_313 @p_313)) :rule refl)
(step t330 (cl (= @p_384 @p_384)) :rule refl)
(step t331 (cl (= @p_119 @p_119)) :rule refl)
(step t332 (cl (! (= @p_272 @p_269) :named @p_432)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_432 1 6))
(step t333 (cl @p_431) :rule cong :premises (t2 t329 t330 t331 t332))
(step t334 (cl (! (=> (! (and @p_79 @p_254 @p_298 @p_117 @p_270) :named @p_388) (! (and @p_270 @p_254 @p_117 @p_79 @p_298) :named @p_390)) :named @p_430) @p_388) :rule implies_neg1)
(anchor :step t335)
(assume t335.a0 @p_79)
(assume t335.a1 @p_254)
(assume t335.a2 @p_298)
(assume t335.a3 @p_117)
(assume t335.a4 @p_270)
(step t335.t0 (cl @p_390 @p_272 @p_313 @p_119 @p_312 @p_384) :rule and_neg)
(step t335.t1 (cl (not (! (= @p_268 (! (and @p_298 (! (not (! (>= @p_295 4) :named @p_296)) :named @p_297)) :named @p_299)) :named @p_300)) (not @p_268) @p_299) :rule equiv_pos2)
(step t335.t2 (cl (! (= @p_267 @p_298) :named @p_311)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_311 3 7))
(step t335.t3 (cl (! (= @p_265 (! (not (! (>= @p_6 @p_264) :named @p_301)) :named @p_302)) :named @p_310)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_310 3 6))
(step t335.t4 (cl (! (= @p_263 (! (+ 1 @p_12) :named @p_306)) :named @p_309)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_309 3 7))
(step t335.t5 (cl (= @p_264 (! (* 4 @p_306) :named @p_307))) :rule cong :premises (t124 t335.t4))
(step t335.t6 (cl (! (= @p_307 (! (+ 4 @p_266) :named @p_303)) :named @p_308)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_308 3 7))
(step t335.t7 (cl (= @p_264 @p_303)) :rule trans :premises (t335.t5 t335.t6))
(step t335.t8 (cl (= @p_301 (! (>= @p_6 @p_303) :named @p_304))) :rule cong :premises (t96 t335.t7))
(step t335.t9 (cl (! (= @p_304 @p_296) :named @p_305)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_305 3 7))
(step t335.t10 (cl (= @p_301 @p_296)) :rule trans :premises (t335.t8 t335.t9))
(step t335.t11 (cl (= @p_302 @p_297)) :rule cong :premises (t335.t10))
(step t335.t12 (cl (= @p_265 @p_297)) :rule trans :premises (t335.t3 t335.t11))
(step t335.t13 (cl @p_300) :rule cong :premises (t335.t2 t335.t12))
(step t335.t14 (cl @p_268) :rule and :premises (t305) :args (1))
(step t335.t15 (cl @p_299) :rule resolution :premises (t335.t1 t335.t13 t335.t14))
(step t335.t16 (cl @p_298) :rule and :premises (t335.t15) :args (0))
(step t335.t17 (cl @p_390) :rule resolution :premises (t335.t0 t335.a4 t335.a1 t335.a3 t335.a0 t335.t16))
(step t335 (cl @p_312 @p_313 @p_384 @p_119 @p_272 @p_390) :rule subproof :discharge (t335.a0 t335.a1 t335.a2 t335.a3 t335.a4))
(step t336 (cl (! (not @p_388) :named @p_389) @p_79) :rule and_pos :args (0))
(step t337 (cl @p_389 @p_254) :rule and_pos :args (1))
(step t338 (cl @p_389 @p_298) :rule and_pos :args (2))
(step t339 (cl @p_389 @p_117) :rule and_pos :args (3))
(step t340 (cl @p_389 @p_270) :rule and_pos :args (4))
(step t341 (cl @p_390 @p_389 @p_389 @p_389 @p_389 @p_389) :rule resolution :premises (t335 t336 t337 t338 t339 t340))
(step t342 (cl @p_389 @p_389 @p_389 @p_389 @p_389 @p_390) :rule reordering :premises (t341))
(step t343 (cl @p_389 @p_390) :rule contraction :premises (t342))
(step t344 (cl @p_430 @p_390) :rule resolution :premises (t334 t343))
(step t345 (cl @p_430 (! (not @p_390) :named @p_391)) :rule implies_neg2)
(step t346 (cl @p_430 @p_430) :rule resolution :premises (t344 t345))
(step t347 (cl @p_430) :rule contraction :premises (t346))
(step t348 (cl @p_389 @p_390) :rule implies :premises (t347))
(step t349 (cl (! (=> @p_390 false) :named @p_392) @p_390) :rule implies_neg1)
(anchor :step t350)
(assume t350.a0 @p_270)
(assume t350.a1 @p_254)
(assume t350.a2 @p_117)
(assume t350.a3 @p_79)
(assume t350.a4 @p_298)
(step t350.t0 (cl (not @p_300) (not @p_268) @p_299) :rule equiv_pos2)
(step t350.t1 (cl @p_311) :rule hole :args ("TRUST_THEORY_REWRITE" @p_311 3 7))
(step t350.t2 (cl @p_310) :rule hole :args ("TRUST_THEORY_REWRITE" @p_310 3 6))
(step t350.t3 (cl @p_309) :rule hole :args ("TRUST_THEORY_REWRITE" @p_309 3 7))
(step t350.t4 (cl (= @p_264 @p_307)) :rule cong :premises (t124 t350.t3))
(step t350.t5 (cl @p_308) :rule hole :args ("TRUST_THEORY_REWRITE" @p_308 3 7))
(step t350.t6 (cl (= @p_264 @p_303)) :rule trans :premises (t350.t4 t350.t5))
(step t350.t7 (cl (= @p_301 @p_304)) :rule cong :premises (t96 t350.t6))
(step t350.t8 (cl @p_305) :rule hole :args ("TRUST_THEORY_REWRITE" @p_305 3 7))
(step t350.t9 (cl (= @p_301 @p_296)) :rule trans :premises (t350.t7 t350.t8))
(step t350.t10 (cl (= @p_302 @p_297)) :rule cong :premises (t350.t9))
(step t350.t11 (cl (= @p_265 @p_297)) :rule trans :premises (t350.t2 t350.t10))
(step t350.t12 (cl @p_300) :rule cong :premises (t350.t1 t350.t11))
(step t350.t13 (cl @p_268) :rule and :premises (t305) :args (1))
(step t350.t14 (cl @p_299) :rule resolution :premises (t350.t0 t350.t12 t350.t13))
(step t350.t15 (cl @p_298) :rule and :premises (t350.t14) :args (0))
(step t350.t16 (cl (not (! (= (! (< @p_295 0) :named @p_393) @p_384) :named @p_429)) (not @p_393) @p_384) :rule equiv_pos2)
(step t350.t17 (cl @p_429) :rule hole :args ("TRUST_THEORY_REWRITE" @p_429 3 6))
(step t350.t18 (cl (not (! (= @p_384 @p_393) :named @p_428)) (! (not @p_384) :named @p_387) @p_393) :rule equiv_pos2)
(step t350.t19 (cl @p_428) :rule symm :premises (t350.t17))
(step t350.t20 (cl (! (=> @p_298 false) :named @p_394) @p_298) :rule implies_neg1)
(anchor :step t350.t21)
(assume t350.t21.a0 @p_298)
(step t350.t21.t0 (cl (not (! (= (! (< (! (+ (! (* -1 @p_295) :named @p_400) (! (* 4 @p_334) :named @p_399) (! (* -4 @p_17) :named @p_372) (! (* 4 (! (+ @p_14 (* -1 @p_12)) :named @p_330)) :named @p_398) @p_6) :named @p_401) (! (+ @p_136 (! (* 4 0) :named @p_395) (! (* -4 1) :named @p_396) @p_395 4) :named @p_397)) :named @p_402) false) :named @p_412)) (not @p_402) false) :rule equiv_pos2)
(step t350.t21.t1 (cl (! (= @p_402 (! (not (! (>= @p_401 @p_397) :named @p_413)) :named @p_414)) :named @p_427)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_427 3 6))
(step t350.t21.t2 (cl (! (= @p_400 (! (+ @p_140 @p_266) :named @p_421)) :named @p_426)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_426 3 7))
(step t350.t21.t3 (cl (! (= @p_399 (! (+ (! (* 4 @p_17) :named @p_332) (! (* -4 @p_14) :named @p_373)) :named @p_420)) :named @p_425)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_425 3 7))
(step t350.t21.t4 (cl (= @p_372 @p_372)) :rule refl)
(step t350.t21.t5 (cl (! (= @p_398 (! (+ (! (* 4 @p_14) :named @p_374) @p_294) :named @p_419)) :named @p_424)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_424 3 7))
(step t350.t21.t6 (cl (= @p_401 (! (+ @p_421 @p_420 @p_372 @p_419 @p_6) :named @p_422))) :rule cong :premises (t350.t21.t2 t350.t21.t3 t350.t21.t4 t350.t21.t5 t96))
(step t350.t21.t7 (cl (! (= @p_422 0) :named @p_423)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_423 3 7))
(step t350.t21.t8 (cl (= @p_401 0)) :rule trans :premises (t350.t21.t6 t350.t21.t7))
(step t350.t21.t9 (cl @p_199) :rule hole :args ("TRUST_THEORY_REWRITE" @p_199 3 6))
(step t350.t21.t10 (cl (! (= @p_395 0) :named @p_417)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_417 3 6))
(step t350.t21.t11 (cl (! (= @p_396 -4) :named @p_418)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_418 3 7))
(step t350.t21.t12 (cl (= @p_397 (! (+ 0 0 -4 0 4) :named @p_415))) :rule cong :premises (t350.t21.t9 t350.t21.t10 t350.t21.t11 t350.t21.t10 t124))
(step t350.t21.t13 (cl (! (= @p_415 0) :named @p_416)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_416 3 7))
(step t350.t21.t14 (cl (= @p_397 0)) :rule trans :premises (t350.t21.t12 t350.t21.t13))
(step t350.t21.t15 (cl (= @p_413 @p_195)) :rule cong :premises (t350.t21.t8 t350.t21.t14))
(step t350.t21.t16 (cl @p_196) :rule hole :args ("TRUST_THEORY_REWRITE" @p_196 3 7))
(step t350.t21.t17 (cl (= @p_413 true)) :rule trans :premises (t350.t21.t15 t350.t21.t16))
(step t350.t21.t18 (cl (= @p_414 @p_193)) :rule cong :premises (t350.t21.t17))
(step t350.t21.t19 (cl @p_194) :rule hole :args ("TRUST_THEORY_REWRITE" @p_194 1 7))
(step t350.t21.t20 (cl (= @p_414 false)) :rule trans :premises (t350.t21.t18 t350.t21.t19))
(step t350.t21.t21 (cl @p_412) :rule trans :premises (t350.t21.t1 t350.t21.t20))
(step t350.t21.t22 (cl (not (! (<= @p_400 @p_136) :named @p_410)) (not (! (= @p_399 @p_395) :named @p_408)) (not (! (<= @p_372 @p_396) :named @p_406)) (not (! (= @p_398 @p_395) :named @p_404)) (! (not (! (< @p_6 4) :named @p_274)) :named @p_293) @p_402) :rule la_generic :args (1/1 -1/1 1/1 -1/1 1/1 1/1))
(step t350.t21.t23 (cl (=> (! (and @p_147 @p_298) :named @p_411) @p_410)) :rule la_mult_neg)
(step t350.t21.t24 (cl (not @p_411) @p_410) :rule implies :premises (t350.t21.t23))
(step t350.t21.t25 (cl @p_411 @p_150 @p_384) :rule and_neg)
(step t350.t21.t26 (cl (= @p_149 @p_147)) :rule equiv_simplify)
(step t350.t21.t27 (cl (not @p_149) @p_147) :rule equiv1 :premises (t350.t21.t26))
(step t350.t21.t28 (cl @p_149) :rule rare_rewrite :args ("evaluate"))
(step t350.t21.t29 (cl @p_147) :rule resolution :premises (t350.t21.t27 t350.t21.t28))
(step t350.t21.t30 (cl @p_411) :rule resolution :premises (t350.t21.t25 t350.t21.t29 t350.t15))
(step t350.t21.t31 (cl @p_410) :rule resolution :premises (t350.t21.t24 t350.t21.t30))
(step t350.t21.t32 (cl (=> (! (and (! (> 4 0) :named @p_350) @p_355) :named @p_409) @p_408)) :rule la_mult_pos)
(step t350.t21.t33 (cl (not @p_409) @p_408) :rule implies :premises (t350.t21.t32))
(step t350.t21.t34 (cl @p_409 (! (not @p_350) :named @p_353) @p_361) :rule and_neg)
(step t350.t21.t35 (cl (= (! (= @p_350 true) :named @p_352) @p_350)) :rule equiv_simplify)
(step t350.t21.t36 (cl (not @p_352) @p_350) :rule equiv1 :premises (t350.t21.t35))
(step t350.t21.t37 (cl @p_352) :rule rare_rewrite :args ("evaluate"))
(step t350.t21.t38 (cl @p_350) :rule resolution :premises (t350.t21.t36 t350.t21.t37))
(step t350.t21.t39 (cl @p_360 @p_312 @p_355) :rule equiv_pos2)
(step t350.t21.t40 (cl @p_359) :rule hole :args ("TRUST_THEORY_REWRITE" @p_359 3 7))
(step t350.t21.t41 (cl @p_358) :rule symm :premises (t350.t21.t40))
(step t350.t21.t42 (cl @p_355) :rule resolution :premises (t350.t21.t39 t350.t21.t41 t350.a3))
(step t350.t21.t43 (cl @p_409) :rule resolution :premises (t350.t21.t34 t350.t21.t38 t350.t21.t42))
(step t350.t21.t44 (cl @p_408) :rule resolution :premises (t350.t21.t33 t350.t21.t43))
(step t350.t21.t45 (cl (=> (! (and (! (< -4 0) :named @p_340) @p_117) :named @p_407) @p_406)) :rule la_mult_neg)
(step t350.t21.t46 (cl (not @p_407) @p_406) :rule implies :premises (t350.t21.t45))
(step t350.t21.t47 (cl @p_407 (! (not @p_340) :named @p_347) @p_119) :rule and_neg)
(step t350.t21.t48 (cl (= (! (= @p_340 true) :named @p_346) @p_340)) :rule equiv_simplify)
(step t350.t21.t49 (cl (not @p_346) @p_340) :rule equiv1 :premises (t350.t21.t48))
(step t350.t21.t50 (cl @p_346) :rule rare_rewrite :args ("evaluate"))
(step t350.t21.t51 (cl @p_340) :rule resolution :premises (t350.t21.t49 t350.t21.t50))
(step t350.t21.t52 (cl @p_407) :rule resolution :premises (t350.t21.t47 t350.t21.t51 t350.a2))
(step t350.t21.t53 (cl @p_406) :rule resolution :premises (t350.t21.t46 t350.t21.t52))
(step t350.t21.t54 (cl (=> (! (and @p_350 (! (= @p_330 0) :named @p_339)) :named @p_405) @p_404)) :rule la_mult_pos)
(step t350.t21.t55 (cl (not @p_405) @p_404) :rule implies :premises (t350.t21.t54))
(step t350.t21.t56 (cl @p_405 @p_353 (! (not @p_339) :named @p_348)) :rule and_neg)
(step t350.t21.t57 (cl (! (not (! (= @p_254 @p_339) :named @p_343)) :named @p_345) @p_313 @p_339) :rule equiv_pos2)
(step t350.t21.t58 (cl (! (= @p_339 @p_254) :named @p_344)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_344 3 7))
(step t350.t21.t59 (cl @p_343) :rule symm :premises (t350.t21.t58))
(step t350.t21.t60 (cl @p_339) :rule resolution :premises (t350.t21.t57 t350.t21.t59 t350.a1))
(step t350.t21.t61 (cl @p_405) :rule resolution :premises (t350.t21.t56 t350.t21.t38 t350.t21.t60))
(step t350.t21.t62 (cl @p_404) :rule resolution :premises (t350.t21.t55 t350.t21.t61))
(step t350.t21.t63 (cl (! (not (! (= @p_270 @p_274) :named @p_290)) :named @p_292) @p_272 @p_274) :rule equiv_pos2)
(step t350.t21.t64 (cl (! (= @p_274 @p_270) :named @p_291)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_291 3 6))
(step t350.t21.t65 (cl @p_290) :rule symm :premises (t350.t21.t64))
(step t350.t21.t66 (cl @p_274) :rule resolution :premises (t350.t21.t63 t350.t21.t65 t350.a0))
(step t350.t21.t67 (cl @p_402) :rule resolution :premises (t350.t21.t22 t350.t21.t31 t350.t21.t44 t350.t21.t53 t350.t21.t62 t350.t21.t66))
(step t350.t21.t68 (cl false) :rule resolution :premises (t350.t21.t0 t350.t21.t21 t350.t21.t67))
(step t350.t21 (cl @p_384 false) :rule subproof :discharge (t350.t21.a0))
(step t350.t22 (cl @p_394 false) :rule resolution :premises (t350.t20 t350.t21))
(step t350.t23 (cl @p_394 @p_132) :rule implies_neg2)
(step t350.t24 (cl @p_394 @p_394) :rule resolution :premises (t350.t22 t350.t23))
(step t350.t25 (cl @p_394) :rule contraction :premises (t350.t24))
(step t350.t26 (cl (= @p_394 @p_384)) :rule implies_simplify)
(step t350.t27 (cl (not @p_394) @p_384) :rule equiv1 :premises (t350.t26))
(step t350.t28 (cl @p_384) :rule resolution :premises (t350.t25 t350.t27))
(step t350.t29 (cl @p_393) :rule resolution :premises (t350.t18 t350.t19 t350.t28))
(step t350.t30 (cl @p_384) :rule resolution :premises (t350.t16 t350.t17 t350.t29))
(step t350.t31 (cl) :rule resolution :premises (t350.t15 t350.t30))
(step t350 (cl @p_272 @p_313 @p_119 @p_312 @p_384 false) :rule subproof :discharge (t350.a0 t350.a1 t350.a2 t350.a3 t350.a4))
(step t351 (cl @p_391 @p_270) :rule and_pos :args (0))
(step t352 (cl @p_391 @p_254) :rule and_pos :args (1))
(step t353 (cl @p_391 @p_117) :rule and_pos :args (2))
(step t354 (cl @p_391 @p_79) :rule and_pos :args (3))
(step t355 (cl @p_391 @p_298) :rule and_pos :args (4))
(step t356 (cl false @p_391 @p_391 @p_391 @p_391 @p_391) :rule resolution :premises (t350 t351 t352 t353 t354 t355))
(step t357 (cl @p_391 @p_391 @p_391 @p_391 @p_391 false) :rule reordering :premises (t356))
(step t358 (cl @p_391 false) :rule contraction :premises (t357))
(step t359 (cl @p_392 false) :rule resolution :premises (t349 t358))
(step t360 (cl @p_392 @p_132) :rule implies_neg2)
(step t361 (cl @p_392 @p_392) :rule resolution :premises (t359 t360))
(step t362 (cl @p_392) :rule contraction :premises (t361))
(step t363 (cl (= @p_392 @p_391)) :rule implies_simplify)
(step t364 (cl (not @p_392) @p_391) :rule equiv1 :premises (t363))
(step t365 (cl @p_391) :rule resolution :premises (t362 t364))
(step t366 (cl @p_389) :rule resolution :premises (t348 t365))
(step t367 (cl @p_312 @p_313 @p_384 @p_119 @p_272) :rule not_and :premises (t366))
(step t368 (cl @p_386 @p_318) :rule or_neg :args (0))
(step t369 (cl @p_386 (! (not @p_313) :named @p_317)) :rule or_neg :args (1))
(step t370 (cl @p_386 @p_387) :rule or_neg :args (2))
(step t371 (cl @p_386 @p_120) :rule or_neg :args (3))
(step t372 (cl @p_386 (not @p_272)) :rule or_neg :args (4))
(step t373 (cl @p_386 @p_386 @p_386 @p_386 @p_386) :rule resolution :premises (t367 t368 t369 t370 t371 t372))
(step t374 (cl @p_386) :rule contraction :premises (t373))
(step t375 (cl @p_385) :rule resolution :premises (t328 t333 t374))
(step t376 (cl @p_312 @p_313 @p_384 @p_119 @p_269) :rule or :premises (t375))
(step t377 (cl (not @p_300) (not @p_268) @p_299) :rule equiv_pos2)
(step t378 (cl @p_311) :rule hole :args ("TRUST_THEORY_REWRITE" @p_311 3 7))
(step t379 (cl @p_310) :rule hole :args ("TRUST_THEORY_REWRITE" @p_310 3 6))
(step t380 (cl @p_309) :rule hole :args ("TRUST_THEORY_REWRITE" @p_309 3 7))
(step t381 (cl (= @p_264 @p_307)) :rule cong :premises (t124 t380))
(step t382 (cl @p_308) :rule hole :args ("TRUST_THEORY_REWRITE" @p_308 3 7))
(step t383 (cl (= @p_264 @p_303)) :rule trans :premises (t381 t382))
(step t384 (cl (= @p_301 @p_304)) :rule cong :premises (t96 t383))
(step t385 (cl @p_305) :rule hole :args ("TRUST_THEORY_REWRITE" @p_305 3 7))
(step t386 (cl (= @p_301 @p_296)) :rule trans :premises (t384 t385))
(step t387 (cl (= @p_302 @p_297)) :rule cong :premises (t386))
(step t388 (cl (= @p_265 @p_297)) :rule trans :premises (t379 t387))
(step t389 (cl @p_300) :rule cong :premises (t378 t388))
(step t390 (cl @p_268) :rule and :premises (t305) :args (1))
(step t391 (cl @p_299) :rule resolution :premises (t377 t389 t390))
(step t392 (cl @p_298) :rule and :premises (t391) :args (0))
(step t393 (cl (not (! (= (! (or @p_115 @p_122 @p_312 @p_313 (! (not @p_297) :named @p_315) @p_120) :named @p_316) (! (or @p_115 @p_18 @p_312 @p_313 @p_296 @p_117) :named @p_314)) :named @p_383)) (not @p_316) @p_314) :rule equiv_pos2)
(step t394 (cl (! (= @p_315 @p_296) :named @p_380)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_380 1 6))
(step t395 (cl @p_383) :rule cong :premises (t214 t1 t2 t329 t394 t217))
(step t396 (cl (! (=> (! (and @p_96 @p_19 @p_79 @p_254 @p_297 @p_119) :named @p_319) (! (and @p_96 @p_254 @p_19 @p_119 @p_79 @p_297) :named @p_321)) :named @p_382) @p_319) :rule implies_neg1)
(anchor :step t397)
(assume t397.a0 @p_96)
(assume t397.a1 @p_19)
(assume t397.a2 @p_79)
(assume t397.a3 @p_254)
(assume t397.a4 @p_297)
(assume t397.a5 @p_119)
(step t397.t0 (cl @p_321 @p_115 @p_313 @p_122 @p_120 @p_312 @p_315) :rule and_neg)
(step t397.t1 (cl @p_297) :rule and :premises (t391) :args (1))
(step t397.t2 (cl @p_321) :rule resolution :premises (t397.t0 t265 t397.a3 t298 t397.a5 t397.a2 t397.t1))
(step t397 (cl @p_115 @p_122 @p_312 @p_313 @p_315 @p_120 @p_321) :rule subproof :discharge (t397.a0 t397.a1 t397.a2 t397.a3 t397.a4 t397.a5))
(step t398 (cl (! (not @p_319) :named @p_320) @p_96) :rule and_pos :args (0))
(step t399 (cl @p_320 @p_19) :rule and_pos :args (1))
(step t400 (cl @p_320 @p_79) :rule and_pos :args (2))
(step t401 (cl @p_320 @p_254) :rule and_pos :args (3))
(step t402 (cl @p_320 @p_297) :rule and_pos :args (4))
(step t403 (cl @p_320 @p_119) :rule and_pos :args (5))
(step t404 (cl @p_321 @p_320 @p_320 @p_320 @p_320 @p_320 @p_320) :rule resolution :premises (t397 t398 t399 t400 t401 t402 t403))
(step t405 (cl @p_320 @p_320 @p_320 @p_320 @p_320 @p_320 @p_321) :rule reordering :premises (t404))
(step t406 (cl @p_320 @p_321) :rule contraction :premises (t405))
(step t407 (cl @p_382 @p_321) :rule resolution :premises (t396 t406))
(step t408 (cl @p_382 (! (not @p_321) :named @p_322)) :rule implies_neg2)
(step t409 (cl @p_382 @p_382) :rule resolution :premises (t407 t408))
(step t410 (cl @p_382) :rule contraction :premises (t409))
(step t411 (cl @p_320 @p_321) :rule implies :premises (t410))
(step t412 (cl (! (=> @p_321 false) :named @p_323) @p_321) :rule implies_neg1)
(anchor :step t413)
(assume t413.a0 @p_96)
(assume t413.a1 @p_254)
(assume t413.a2 @p_19)
(assume t413.a3 @p_119)
(assume t413.a4 @p_79)
(assume t413.a5 @p_297)
(step t413.t0 (cl (! (not (! (= @p_297 (! (< @p_295 4) :named @p_324)) :named @p_362)) :named @p_364) @p_315 @p_324) :rule equiv_pos2)
(step t413.t1 (cl (! (= @p_324 @p_297) :named @p_363)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_363 3 6))
(step t413.t2 (cl @p_362) :rule symm :premises (t413.t1))
(step t413.t3 (cl @p_297) :rule and :premises (t391) :args (1))
(step t413.t4 (cl @p_324) :rule resolution :premises (t413.t0 t413.t2 t413.t3))
(step t413.t5 (cl (not (! (= @p_296 (! (not @p_324) :named @p_325)) :named @p_381)) @p_297 @p_325) :rule equiv_pos2)
(step t413.t6 (cl (= @p_325 @p_315)) :rule cong :premises (t413.t1))
(step t413.t7 (cl (! (= @p_325 @p_296) :named @p_379)) :rule trans :premises (t413.t6 t394))
(step t413.t8 (cl @p_381) :rule symm :premises (t413.t7))
(step t413.t9 (cl (not @p_379) (not @p_325) @p_296) :rule equiv_pos2)
(step t413.t10 (cl (! (=> @p_324 false) :named @p_326) @p_324) :rule implies_neg1)
(anchor :step t413.t11)
(assume t413.t11.a0 @p_324)
(step t413.t11.t0 (cl (not (! (= (! (< (! (+ @p_295 (! (* -4 @p_334) :named @p_335) @p_332 (! (* -4 @p_330) :named @p_331) @p_140) :named @p_336) (! (+ 4 (! (* -4 0) :named @p_327) (! (* 4 -1) :named @p_328) @p_327 @p_136) :named @p_329)) :named @p_337) false) :named @p_365)) (not @p_337) false) :rule equiv_pos2)
(step t413.t11.t1 (cl (! (= @p_337 (! (not (! (>= @p_336 @p_329) :named @p_366)) :named @p_367)) :named @p_378)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_378 3 6))
(step t413.t11.t2 (cl (! (= @p_336 (! (+ @p_332 @p_140 @p_266 @p_374 @p_294 @p_373 @p_372 @p_6) :named @p_375)) :named @p_377)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_377 3 6))
(step t413.t11.t3 (cl (! (= @p_375 0) :named @p_376)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_376 3 7))
(step t413.t11.t4 (cl (= @p_336 0)) :rule trans :premises (t413.t11.t2 t413.t11.t3))
(step t413.t11.t5 (cl (! (= @p_327 0) :named @p_370)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_370 3 6))
(step t413.t11.t6 (cl (! (= @p_328 -4) :named @p_371)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_371 3 7))
(step t413.t11.t7 (cl @p_199) :rule hole :args ("TRUST_THEORY_REWRITE" @p_199 3 6))
(step t413.t11.t8 (cl (= @p_329 (! (+ 4 0 -4 0 0) :named @p_368))) :rule cong :premises (t124 t413.t11.t5 t413.t11.t6 t413.t11.t5 t413.t11.t7))
(step t413.t11.t9 (cl (! (= @p_368 0) :named @p_369)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_369 3 7))
(step t413.t11.t10 (cl (= @p_329 0)) :rule trans :premises (t413.t11.t8 t413.t11.t9))
(step t413.t11.t11 (cl (= @p_366 @p_195)) :rule cong :premises (t413.t11.t4 t413.t11.t10))
(step t413.t11.t12 (cl @p_196) :rule hole :args ("TRUST_THEORY_REWRITE" @p_196 3 7))
(step t413.t11.t13 (cl (= @p_366 true)) :rule trans :premises (t413.t11.t11 t413.t11.t12))
(step t413.t11.t14 (cl (= @p_367 @p_193)) :rule cong :premises (t413.t11.t13))
(step t413.t11.t15 (cl @p_194) :rule hole :args ("TRUST_THEORY_REWRITE" @p_194 1 7))
(step t413.t11.t16 (cl (= @p_367 false)) :rule trans :premises (t413.t11.t14 t413.t11.t15))
(step t413.t11.t17 (cl @p_365) :rule trans :premises (t413.t11.t1 t413.t11.t16))
(step t413.t11.t18 (cl @p_325 (not (! (= @p_335 @p_327) :named @p_354)) (not (! (<= @p_332 @p_328) :named @p_349)) (not (! (= @p_331 @p_327) :named @p_338)) @p_189 @p_337) :rule la_generic :args (1/1 -1/1 1/1 -1/1 1/1 1/1))
(step t413.t11.t19 (cl @p_364 @p_315 @p_324) :rule equiv_pos2)
(step t413.t11.t20 (cl @p_297) :rule and :premises (t391) :args (1))
(step t413.t11.t21 (cl @p_324) :rule resolution :premises (t413.t11.t19 t413.t2 t413.t11.t20))
(step t413.t11.t22 (cl (=> (! (and @p_340 @p_355) :named @p_356) @p_354)) :rule la_mult_neg)
(step t413.t11.t23 (cl (not @p_356) @p_354) :rule implies :premises (t413.t11.t22))
(step t413.t11.t24 (cl @p_356 @p_347 @p_361) :rule and_neg)
(step t413.t11.t25 (cl (= @p_346 @p_340)) :rule equiv_simplify)
(step t413.t11.t26 (cl (not @p_346) @p_340) :rule equiv1 :premises (t413.t11.t25))
(step t413.t11.t27 (cl @p_346) :rule rare_rewrite :args ("evaluate"))
(step t413.t11.t28 (cl @p_340) :rule resolution :premises (t413.t11.t26 t413.t11.t27))
(step t413.t11.t29 (cl @p_360 @p_312 @p_355) :rule equiv_pos2)
(step t413.t11.t30 (cl @p_359) :rule hole :args ("TRUST_THEORY_REWRITE" @p_359 3 7))
(step t413.t11.t31 (cl @p_358) :rule symm :premises (t413.t11.t30))
(step t413.t11.t32 (cl @p_355) :rule resolution :premises (t413.t11.t29 t413.t11.t31 t413.a4))
(step t413.t11.t33 (cl @p_356) :rule resolution :premises (t413.t11.t24 t413.t11.t28 t413.t11.t32))
(step t413.t11.t34 (cl @p_354) :rule resolution :premises (t413.t11.t23 t413.t11.t33))
(step t413.t11.t35 (cl (=> (! (and @p_350 @p_152) :named @p_351) @p_349)) :rule la_mult_pos)
(step t413.t11.t36 (cl (not @p_351) @p_349) :rule implies :premises (t413.t11.t35))
(step t413.t11.t37 (cl @p_351 @p_353 @p_174) :rule and_neg)
(step t413.t11.t38 (cl (= @p_352 @p_350)) :rule equiv_simplify)
(step t413.t11.t39 (cl (not @p_352) @p_350) :rule equiv1 :premises (t413.t11.t38))
(step t413.t11.t40 (cl @p_352) :rule rare_rewrite :args ("evaluate"))
(step t413.t11.t41 (cl @p_350) :rule resolution :premises (t413.t11.t39 t413.t11.t40))
(step t413.t11.t42 (cl @p_171 @p_152) :rule la_generic :args (1/1 1/1))
(step t413.t11.t43 (cl @p_170) :rule la_disequality)
(step t413.t11.t44 (cl @p_18 @p_169 @p_157) :rule or :premises (t413.t11.t43))
(step t413.t11.t45 (cl @p_168 @p_161) :rule la_generic :args (1/1 1/1))
(step t413.t11.t46 (cl @p_167 @p_120 @p_162) :rule equiv_pos2)
(step t413.t11.t47 (cl @p_166) :rule hole :args ("TRUST_THEORY_REWRITE" @p_166 3 6))
(step t413.t11.t48 (cl @p_165) :rule symm :premises (t413.t11.t47))
(step t413.t11.t49 (cl @p_162) :rule resolution :premises (t413.t11.t46 t413.t11.t48 t413.a3))
(step t413.t11.t50 (cl @p_161) :rule resolution :premises (t413.t11.t45 t413.t11.t49))
(step t413.t11.t51 (cl @p_159 @p_155 @p_160) :rule equiv_pos1)
(step t413.t11.t52 (cl @p_158) :rule comp_simplify)
(step t413.t11.t53 (cl @p_155) :rule resolution :premises (t413.t11.t44 t298 t413.t11.t50 t413.t11.t51 t413.t11.t52))
(step t413.t11.t54 (cl @p_152) :rule resolution :premises (t413.t11.t42 t413.t11.t53))
(step t413.t11.t55 (cl @p_351) :rule resolution :premises (t413.t11.t37 t413.t11.t41 t413.t11.t54))
(step t413.t11.t56 (cl @p_349) :rule resolution :premises (t413.t11.t36 t413.t11.t55))
(step t413.t11.t57 (cl (=> (! (and @p_340 @p_339) :named @p_341) @p_338)) :rule la_mult_neg)
(step t413.t11.t58 (cl (not @p_341) @p_338) :rule implies :premises (t413.t11.t57))
(step t413.t11.t59 (cl @p_341 @p_347 @p_348) :rule and_neg)
(step t413.t11.t60 (cl @p_345 @p_313 @p_339) :rule equiv_pos2)
(step t413.t11.t61 (cl @p_344) :rule hole :args ("TRUST_THEORY_REWRITE" @p_344 3 7))
(step t413.t11.t62 (cl @p_343) :rule symm :premises (t413.t11.t61))
(step t413.t11.t63 (cl @p_339) :rule resolution :premises (t413.t11.t60 t413.t11.t62 t413.a1))
(step t413.t11.t64 (cl @p_341) :rule resolution :premises (t413.t11.t59 t413.t11.t28 t413.t11.t63))
(step t413.t11.t65 (cl @p_338) :rule resolution :premises (t413.t11.t58 t413.t11.t64))
(step t413.t11.t66 (cl (=> @p_148 @p_146)) :rule la_mult_neg)
(step t413.t11.t67 (cl (not @p_148) @p_146) :rule implies :premises (t413.t11.t66))
(step t413.t11.t68 (cl @p_148 @p_150 @p_115) :rule and_neg)
(step t413.t11.t69 (cl (= @p_149 @p_147)) :rule equiv_simplify)
(step t413.t11.t70 (cl (not @p_149) @p_147) :rule equiv1 :premises (t413.t11.t69))
(step t413.t11.t71 (cl @p_149) :rule rare_rewrite :args ("evaluate"))
(step t413.t11.t72 (cl @p_147) :rule resolution :premises (t413.t11.t70 t413.t11.t71))
(step t413.t11.t73 (cl @p_148) :rule resolution :premises (t413.t11.t68 t413.t11.t72 t265))
(step t413.t11.t74 (cl @p_146) :rule resolution :premises (t413.t11.t67 t413.t11.t73))
(step t413.t11.t75 (cl @p_337) :rule resolution :premises (t413.t11.t18 t413.t11.t21 t413.t11.t34 t413.t11.t56 t413.t11.t65 t413.t11.t74))
(step t413.t11.t76 (cl false) :rule resolution :premises (t413.t11.t0 t413.t11.t17 t413.t11.t75))
(step t413.t11 (cl @p_325 false) :rule subproof :discharge (t413.t11.a0))
(step t413.t12 (cl @p_326 false) :rule resolution :premises (t413.t10 t413.t11))
(step t413.t13 (cl @p_326 @p_132) :rule implies_neg2)
(step t413.t14 (cl @p_326 @p_326) :rule resolution :premises (t413.t12 t413.t13))
(step t413.t15 (cl @p_326) :rule contraction :premises (t413.t14))
(step t413.t16 (cl (= @p_326 @p_325)) :rule implies_simplify)
(step t413.t17 (cl (not @p_326) @p_325) :rule equiv1 :premises (t413.t16))
(step t413.t18 (cl @p_325) :rule resolution :premises (t413.t15 t413.t17))
(step t413.t19 (cl @p_296) :rule resolution :premises (t413.t9 t413.t7 t413.t18))
(step t413.t20 (cl @p_325) :rule resolution :premises (t413.t5 t413.t8 t413.t19))
(step t413.t21 (cl) :rule resolution :premises (t413.t4 t413.t20))
(step t413 (cl @p_115 @p_313 @p_122 @p_120 @p_312 @p_315 false) :rule subproof :discharge (t413.a0 t413.a1 t413.a2 t413.a3 t413.a4 t413.a5))
(step t414 (cl @p_322 @p_96) :rule and_pos :args (0))
(step t415 (cl @p_322 @p_254) :rule and_pos :args (1))
(step t416 (cl @p_322 @p_19) :rule and_pos :args (2))
(step t417 (cl @p_322 @p_119) :rule and_pos :args (3))
(step t418 (cl @p_322 @p_79) :rule and_pos :args (4))
(step t419 (cl @p_322 @p_297) :rule and_pos :args (5))
(step t420 (cl false @p_322 @p_322 @p_322 @p_322 @p_322 @p_322) :rule resolution :premises (t413 t414 t415 t416 t417 t418 t419))
(step t421 (cl @p_322 @p_322 @p_322 @p_322 @p_322 @p_322 false) :rule reordering :premises (t420))
(step t422 (cl @p_322 false) :rule contraction :premises (t421))
(step t423 (cl @p_323 false) :rule resolution :premises (t412 t422))
(step t424 (cl @p_323 @p_132) :rule implies_neg2)
(step t425 (cl @p_323 @p_323) :rule resolution :premises (t423 t424))
(step t426 (cl @p_323) :rule contraction :premises (t425))
(step t427 (cl (= @p_323 @p_322)) :rule implies_simplify)
(step t428 (cl (not @p_323) @p_322) :rule equiv1 :premises (t427))
(step t429 (cl @p_322) :rule resolution :premises (t426 t428))
(step t430 (cl @p_320) :rule resolution :premises (t411 t429))
(step t431 (cl @p_115 @p_122 @p_312 @p_313 @p_315 @p_120) :rule not_and :premises (t430))
(step t432 (cl @p_316 @p_126) :rule or_neg :args (0))
(step t433 (cl @p_316 @p_125) :rule or_neg :args (1))
(step t434 (cl @p_316 @p_318) :rule or_neg :args (2))
(step t435 (cl @p_316 @p_317) :rule or_neg :args (3))
(step t436 (cl @p_316 (not @p_315)) :rule or_neg :args (4))
(step t437 (cl @p_316 @p_124) :rule or_neg :args (5))
(step t438 (cl @p_316 @p_316 @p_316 @p_316 @p_316 @p_316) :rule resolution :premises (t431 t432 t433 t434 t435 t436 t437))
(step t439 (cl @p_316) :rule contraction :premises (t438))
(step t440 (cl @p_314) :rule resolution :premises (t393 t395 t439))
(step t441 (cl @p_115 @p_18 @p_312 @p_313 @p_296 @p_117) :rule or :premises (t440))
(step t442 (cl @p_18 @p_296 @p_115 @p_312 @p_313 @p_117) :rule reordering :premises (t441))
(step t443 (cl @p_297) :rule and :premises (t391) :args (1))
(step t444 (cl (! (=> @p_13 @p_270) :named @p_271) @p_13) :rule implies_neg1)
(anchor :step t445)
(assume t445.a0 @p_13)
(step t445.t0 (cl @p_271 @p_13) :rule implies_neg1)
(anchor :step t445.t1)
(assume t445.t1.a0 @p_13)
(step t445.t1.t0 (cl (not @p_291) @p_293 @p_270) :rule equiv_pos2)
(step t445.t1.t1 (cl @p_291) :rule hole :args ("TRUST_THEORY_REWRITE" @p_291 3 6))
(step t445.t1.t2 (cl @p_292 @p_272 @p_274) :rule equiv_pos2)
(step t445.t1.t3 (cl @p_290) :rule symm :premises (t445.t1.t1))
(step t445.t1.t4 (cl (! (=> @p_269 false) :named @p_275) @p_269) :rule implies_neg1)
(anchor :step t445.t1.t5)
(assume t445.t1.t5.a0 @p_269)
(step t445.t1.t5.t0 (cl (not (! (= (! (<= (! (+ @p_140 @p_6) :named @p_278) (! (+ (! (* -1 4) :named @p_276) 2) :named @p_277)) :named @p_279) false) :named @p_283)) (not @p_279) false) :rule equiv_pos2)
(step t445.t1.t5.t1 (cl (! (= @p_278 0) :named @p_289)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_289 3 7))
(step t445.t1.t5.t2 (cl (! (= @p_276 -4) :named @p_288)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_288 3 7))
(step t445.t1.t5.t3 (cl (= @p_277 (! (+ -4 2) :named @p_286))) :rule cong :premises (t445.t1.t5.t2 t97))
(step t445.t1.t5.t4 (cl (! (= @p_286 -2) :named @p_287)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_287 3 7))
(step t445.t1.t5.t5 (cl (= @p_277 -2)) :rule trans :premises (t445.t1.t5.t3 t445.t1.t5.t4))
(step t445.t1.t5.t6 (cl (= @p_279 (! (<= 0 -2) :named @p_284))) :rule cong :premises (t445.t1.t5.t1 t445.t1.t5.t5))
(step t445.t1.t5.t7 (cl (! (= @p_284 false) :named @p_285)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_285 3 7))
(step t445.t1.t5.t8 (cl @p_283) :rule trans :premises (t445.t1.t5.t6 t445.t1.t5.t7))
(step t445.t1.t5.t9 (cl (not (! (<= @p_140 @p_276) :named @p_280)) (! (not @p_13) :named @p_253) @p_279) :rule la_generic :args (1/1 -1/1 1/1))
(step t445.t1.t5.t10 (cl (=> (! (and @p_147 @p_269) :named @p_281) @p_280)) :rule la_mult_neg)
(step t445.t1.t5.t11 (cl (not @p_281) @p_280) :rule implies :premises (t445.t1.t5.t10))
(step t445.t1.t5.t12 (cl @p_281 @p_150 @p_270) :rule and_neg)
(step t445.t1.t5.t13 (cl (= @p_149 @p_147)) :rule equiv_simplify)
(step t445.t1.t5.t14 (cl (not @p_149) @p_147) :rule equiv1 :premises (t445.t1.t5.t13))
(step t445.t1.t5.t15 (cl @p_149) :rule rare_rewrite :args ("evaluate"))
(step t445.t1.t5.t16 (cl @p_147) :rule resolution :premises (t445.t1.t5.t14 t445.t1.t5.t15))
(step t445.t1.t5.t17 (cl @p_281) :rule resolution :premises (t445.t1.t5.t12 t445.t1.t5.t16 t445.t1.t5.a0))
(step t445.t1.t5.t18 (cl @p_280) :rule resolution :premises (t445.t1.t5.t11 t445.t1.t5.t17))
(step t445.t1.t5.t19 (cl @p_279) :rule resolution :premises (t445.t1.t5.t9 t445.t1.t5.t18 t445.t1.a0))
(step t445.t1.t5.t20 (cl false) :rule resolution :premises (t445.t1.t5.t0 t445.t1.t5.t8 t445.t1.t5.t19))
(step t445.t1.t5 (cl @p_270 false) :rule subproof :discharge (t445.t1.t5.a0))
(step t445.t1.t6 (cl @p_275 false) :rule resolution :premises (t445.t1.t4 t445.t1.t5))
(step t445.t1.t7 (cl @p_275 @p_132) :rule implies_neg2)
(step t445.t1.t8 (cl @p_275 @p_275) :rule resolution :premises (t445.t1.t6 t445.t1.t7))
(step t445.t1.t9 (cl @p_275) :rule contraction :premises (t445.t1.t8))
(step t445.t1.t10 (cl (= @p_275 @p_270)) :rule implies_simplify)
(step t445.t1.t11 (cl (not @p_275) @p_270) :rule equiv1 :premises (t445.t1.t10))
(step t445.t1.t12 (cl @p_270) :rule resolution :premises (t445.t1.t9 t445.t1.t11))
(step t445.t1.t13 (cl @p_274) :rule resolution :premises (t445.t1.t2 t445.t1.t3 t445.t1.t12))
(step t445.t1.t14 (cl @p_270) :rule resolution :premises (t445.t1.t0 t445.t1.t1 t445.t1.t13))
(step t445.t1 (cl @p_253 @p_270) :rule subproof :discharge (t445.t1.a0))
(step t445.t2 (cl @p_271 @p_270) :rule resolution :premises (t445.t0 t445.t1))
(step t445.t3 (cl @p_271 @p_272) :rule implies_neg2)
(step t445.t4 (cl @p_271 @p_271) :rule resolution :premises (t445.t2 t445.t3))
(step t445.t5 (cl @p_271) :rule contraction :premises (t445.t4))
(step t445.t6 (cl @p_253 @p_270) :rule implies :premises (t445.t5))
(step t445.t7 (cl @p_270) :rule resolution :premises (t445.t6 t445.a0))
(step t445 (cl @p_253 @p_270) :rule subproof :discharge (t445.a0))
(step t446 (cl @p_271 @p_270) :rule resolution :premises (t444 t445))
(step t447 (cl @p_271 @p_272) :rule implies_neg2)
(step t448 (cl @p_271 @p_271) :rule resolution :premises (t446 t447))
(step t449 (cl @p_271) :rule contraction :premises (t448))
(step t450 (cl @p_253 @p_270) :rule implies :premises (t449))
(step t451 (cl @p_253 @p_254) :rule ite2 :premises (t326))
(step t452 (cl @p_254 @p_253) :rule reordering :premises (t451))
(step t453 (cl @p_253) :rule resolution :premises (t376 t302 t392 t442 t302 t443 t265 t298 t450 t452))
(step t454 (cl @p_252) :rule resolution :premises (t327 t453))
(step t455 (cl (! (not (! (= (! (ite @p_10 @p_433 (! (= @p_11 @p_8) :named @p_434)) :named @p_480) @p_480) :named @p_481)) :named @p_484) (! (not @p_480) :named @p_485) @p_480) :rule equiv_pos2)
(step t456 (cl @p_25) :rule refl)
(step t457 (cl @p_29) :rule and :premises (t108) :args (0))
(step t458 (cl (! (= @p_433 @p_433) :named @p_483)) :rule cong :premises (t456 t457))
(step t459 (cl (= (! (= @p_26 true) :named @p_444) @p_26)) :rule equiv_simplify)
(step t460 (cl (not @p_444) @p_26) :rule equiv1 :premises (t459))
(step t461 (cl (! (= @p_26 @p_26) :named @p_445)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_445 3 7))
(step t462 (cl @p_26) :rule refl)
(step t463 (cl @p_26) :rule trans :premises (t462 t131))
(step t464 (cl @p_445) :rule cong :premises (t463 t131))
(step t465 (cl @p_444) :rule hole :args ("TRUST_THEORY_REWRITE" @p_444 3 6))
(step t466 (cl @p_444) :rule trans :premises (t461 t464 t465))
(step t467 (cl @p_26) :rule resolution :premises (t460 t466))
(step t468 (cl (! (= @p_434 @p_434) :named @p_482)) :rule cong :premises (t456 t467))
(step t469 (cl @p_481) :rule cong :premises (t283 t458 t468))
(step t470 (cl @p_484 @p_485 @p_480) :rule equiv_pos2)
(step t471 (cl @p_483) :rule cong :premises (t316 t284))
(step t472 (cl @p_26) :rule refl)
(step t473 (cl @p_482) :rule cong :premises (t316 t472))
(step t474 (cl @p_481) :rule cong :premises (t283 t471 t473))
(step t475 (cl @p_480) :rule hole)
(step t476 (cl @p_480) :rule resolution :premises (t470 t474 t475))
(step t477 (cl @p_480) :rule resolution :premises (t455 t469 t476))
(step t478 (cl @p_10 @p_434) :rule ite1 :premises (t477))
(step t479 (cl (not (! (= (! (or @p_122 @p_312 @p_486 (! (not (! (= @p_8 0) :named @p_436)) :named @p_504) (! (not @p_434) :named @p_435)) :named @p_506) (! (or @p_18 @p_312 @p_486 @p_504 @p_435) :named @p_505)) :named @p_513)) (not @p_506) @p_505) :rule equiv_pos2)
(step t480 (cl (= @p_504 @p_504)) :rule refl)
(step t481 (cl (= @p_435 @p_435)) :rule refl)
(step t482 (cl @p_513) :rule cong :premises (t1 t2 t3 t480 t481))
(step t483 (cl (! (and @p_19 @p_79 @p_252 @p_436) :named @p_507) @p_122 @p_312 @p_486 @p_504) :rule and_neg)
(step t484 (cl (! (=> @p_507 @p_435) :named @p_509) @p_507) :rule implies_neg1)
(anchor :step t485)
(assume t485.a0 @p_19)
(assume t485.a1 @p_79)
(assume t485.a2 @p_252)
(assume t485.a3 @p_436)
(step t485.t0 (cl @p_509 @p_507) :rule implies_neg1)
(anchor :step t485.t1)
(assume t485.t1.a0 @p_19)
(assume t485.t1.a1 @p_79)
(assume t485.t1.a2 @p_252)
(assume t485.t1.a3 @p_436)
(step t485.t1.t0 (cl (= (! (= @p_434 false) :named @p_512) @p_435)) :rule equiv_simplify)
(step t485.t1.t1 (cl (not @p_512) @p_435) :rule equiv1 :premises (t485.t1.t0))
(step t485.t1.t2 (cl (= @p_11 @p_14)) :rule symm :premises (t485.t1.a2))
(step t485.t1.t3 (cl (! (= @p_14 @p_17) :named @p_502)) :rule symm :premises (t485.t1.a1))
(step t485.t1.t4 (cl (= @p_11 @p_17)) :rule trans :premises (t485.t1.t2 t485.t1.t3))
(step t485.t1.t5 (cl (= @p_434 @p_18)) :rule cong :premises (t485.t1.t4 t485.t1.a3))
(step t485.t1.t6 (cl (= (! (= @p_18 false) :named @p_501) @p_19)) :rule equiv_simplify)
(step t485.t1.t7 (cl @p_501 @p_122) :rule equiv2 :premises (t485.t1.t6))
(step t485.t1.t8 (cl @p_125 @p_18) :rule not_not)
(step t485.t1.t9 (cl @p_501 @p_18) :rule resolution :premises (t485.t1.t7 t485.t1.t8))
(step t485.t1.t10 (cl @p_501) :rule resolution :premises (t485.t1.t9 t298))
(step t485.t1.t11 (cl @p_512) :rule trans :premises (t485.t1.t5 t485.t1.t10))
(step t485.t1.t12 (cl @p_435) :rule resolution :premises (t485.t1.t1 t485.t1.t11))
(step t485.t1 (cl @p_122 @p_312 @p_486 @p_504 @p_435) :rule subproof :discharge (t485.t1.a0 t485.t1.a1 t485.t1.a2 t485.t1.a3))
(step t485.t2 (cl (! (not @p_507) :named @p_508) @p_19) :rule and_pos :args (0))
(step t485.t3 (cl @p_508 @p_79) :rule and_pos :args (1))
(step t485.t4 (cl @p_508 @p_252) :rule and_pos :args (2))
(step t485.t5 (cl @p_508 @p_436) :rule and_pos :args (3))
(step t485.t6 (cl @p_435 @p_508 @p_508 @p_508 @p_508) :rule resolution :premises (t485.t1 t485.t2 t485.t3 t485.t4 t485.t5))
(step t485.t7 (cl @p_508 @p_508 @p_508 @p_508 @p_435) :rule reordering :premises (t485.t6))
(step t485.t8 (cl @p_508 @p_435) :rule contraction :premises (t485.t7))
(step t485.t9 (cl @p_509 @p_435) :rule resolution :premises (t485.t0 t485.t8))
(step t485.t10 (cl @p_509 (! (not @p_435) :named @p_491)) :rule implies_neg2)
(step t485.t11 (cl @p_509 @p_509) :rule resolution :premises (t485.t9 t485.t10))
(step t485.t12 (cl @p_509) :rule contraction :premises (t485.t11))
(step t485.t13 (cl @p_508 @p_435) :rule implies :premises (t485.t12))
(step t485.t14 (cl @p_507 @p_122 @p_312 @p_486 @p_504) :rule and_neg)
(step t485.t15 (cl @p_507) :rule resolution :premises (t485.t14 t298 t485.a1 t485.a2 t485.a3))
(step t485.t16 (cl @p_435) :rule resolution :premises (t485.t13 t485.t15))
(step t485 (cl @p_122 @p_312 @p_486 @p_504 @p_435) :rule subproof :discharge (t485.a0 t485.a1 t485.a2 t485.a3))
(step t486 (cl @p_508 @p_19) :rule and_pos :args (0))
(step t487 (cl @p_508 @p_79) :rule and_pos :args (1))
(step t488 (cl @p_508 @p_252) :rule and_pos :args (2))
(step t489 (cl @p_508 @p_436) :rule and_pos :args (3))
(step t490 (cl @p_435 @p_508 @p_508 @p_508 @p_508) :rule resolution :premises (t485 t486 t487 t488 t489))
(step t491 (cl @p_508 @p_508 @p_508 @p_508 @p_435) :rule reordering :premises (t490))
(step t492 (cl @p_508 @p_435) :rule contraction :premises (t491))
(step t493 (cl @p_509 @p_435) :rule resolution :premises (t484 t492))
(step t494 (cl @p_509 @p_491) :rule implies_neg2)
(step t495 (cl @p_509 @p_509) :rule resolution :premises (t493 t494))
(step t496 (cl @p_509) :rule contraction :premises (t495))
(step t497 (cl @p_508 @p_435) :rule implies :premises (t496))
(step t498 (cl @p_122 @p_312 @p_486 @p_504 @p_435) :rule resolution :premises (t483 t497))
(step t499 (cl @p_506 @p_125) :rule or_neg :args (0))
(step t500 (cl @p_506 @p_318) :rule or_neg :args (1))
(step t501 (cl @p_506 @p_490) :rule or_neg :args (2))
(step t502 (cl @p_506 (not @p_504)) :rule or_neg :args (3))
(step t503 (cl @p_506 @p_491) :rule or_neg :args (4))
(step t504 (cl @p_506 @p_506 @p_506 @p_506 @p_506) :rule resolution :premises (t498 t499 t500 t501 t502 t503))
(step t505 (cl @p_506) :rule contraction :premises (t504))
(step t506 (cl @p_505) :rule resolution :premises (t479 t482 t505))
(step t507 (cl @p_18 @p_312 @p_486 @p_504 @p_435) :rule or :premises (t506))
(step t508 (cl @p_18 @p_312 @p_486 @p_435 @p_504) :rule reordering :premises (t507))
(step t509 (cl (not (! (= (! (ite @p_7 (! (= @p_8 @p_6) :named @p_440) @p_436) :named @p_441) (! (ite @p_7 (! (= @p_6 @p_8) :named @p_438) @p_436) :named @p_439)) :named @p_447)) (! (not @p_441) :named @p_446) @p_439) :rule equiv_pos2)
(step t510 (cl (! (= @p_440 @p_438) :named @p_448)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_448 3 7))
(step t511 (cl (! (= @p_436 @p_436) :named @p_443)) :rule refl)
(step t512 (cl @p_447) :rule cong :premises (t285 t510 t511))
(step t513 (cl (not (! (= @p_441 @p_441) :named @p_442)) @p_446 @p_441) :rule equiv_pos2)
(step t514 (cl (= @p_440 @p_440)) :rule cong :premises (t467 t96))
(step t515 (cl @p_443) :rule cong :premises (t467 t129))
(step t516 (cl @p_442) :rule cong :premises (t285 t514 t515))
(step t517 (cl @p_441) :rule hole)
(step t518 (cl @p_441) :rule resolution :premises (t513 t516 t517))
(step t519 (cl @p_439) :rule resolution :premises (t509 t512 t518))
(step t520 (cl @p_7 @p_436) :rule ite1 :premises (t519))
(step t521 (cl (not (! (= (! (or (! (not @p_7) :named @p_437) @p_122 @p_312 @p_435 (! (not @p_438) :named @p_487) @p_486) :named @p_489) (! (or @p_437 @p_18 @p_312 @p_435 @p_487 @p_486) :named @p_488)) :named @p_503)) (not @p_489) @p_488) :rule equiv_pos2)
(step t522 (cl (= @p_437 @p_437)) :rule refl)
(step t523 (cl (= @p_487 @p_487)) :rule refl)
(step t524 (cl @p_503) :rule cong :premises (t522 t1 t2 t481 t523 t3))
(step t525 (cl (! (and @p_7 @p_19 @p_79 @p_434 @p_438) :named @p_492) @p_437 @p_122 @p_312 @p_435 @p_487) :rule and_neg)
(step t526 (cl (! (=> @p_492 @p_486) :named @p_494) @p_492) :rule implies_neg1)
(anchor :step t527)
(assume t527.a0 @p_7)
(assume t527.a1 @p_19)
(assume t527.a2 @p_79)
(assume t527.a3 @p_434)
(assume t527.a4 @p_438)
(step t527.t0 (cl (! (=> (! (and @p_19 @p_79 @p_7 @p_438 @p_434) :named @p_495) @p_486) :named @p_499) @p_495) :rule implies_neg1)
(anchor :step t527.t1)
(assume t527.t1.a0 @p_19)
(assume t527.t1.a1 @p_79)
(assume t527.t1.a2 @p_7)
(assume t527.t1.a3 @p_438)
(assume t527.t1.a4 @p_434)
(step t527.t1.t0 (cl (= (! (= @p_252 false) :named @p_500) @p_486)) :rule equiv_simplify)
(step t527.t1.t1 (cl (not @p_500) @p_486) :rule equiv1 :premises (t527.t1.t0))
(step t527.t1.t2 (cl @p_502) :rule symm :premises (t527.t1.a1))
(step t527.t1.t3 (cl @p_440) :rule symm :premises (t527.t1.a3))
(step t527.t1.t4 (cl (= @p_11 0)) :rule trans :premises (t527.t1.a4 t527.t1.t3 t527.t1.a2))
(step t527.t1.t5 (cl (= @p_252 @p_18)) :rule cong :premises (t527.t1.t2 t527.t1.t4))
(step t527.t1.t6 (cl (= @p_501 @p_19)) :rule equiv_simplify)
(step t527.t1.t7 (cl @p_501 @p_122) :rule equiv2 :premises (t527.t1.t6))
(step t527.t1.t8 (cl @p_125 @p_18) :rule not_not)
(step t527.t1.t9 (cl @p_501 @p_18) :rule resolution :premises (t527.t1.t7 t527.t1.t8))
(step t527.t1.t10 (cl @p_501) :rule resolution :premises (t527.t1.t9 t298))
(step t527.t1.t11 (cl @p_500) :rule trans :premises (t527.t1.t5 t527.t1.t10))
(step t527.t1.t12 (cl @p_486) :rule resolution :premises (t527.t1.t1 t527.t1.t11))
(step t527.t1 (cl @p_122 @p_312 @p_437 @p_487 @p_435 @p_486) :rule subproof :discharge (t527.t1.a0 t527.t1.a1 t527.t1.a2 t527.t1.a3 t527.t1.a4))
(step t527.t2 (cl (! (not @p_495) :named @p_498) @p_19) :rule and_pos :args (0))
(step t527.t3 (cl @p_498 @p_79) :rule and_pos :args (1))
(step t527.t4 (cl @p_498 @p_7) :rule and_pos :args (2))
(step t527.t5 (cl @p_498 @p_438) :rule and_pos :args (3))
(step t527.t6 (cl @p_498 @p_434) :rule and_pos :args (4))
(step t527.t7 (cl @p_486 @p_498 @p_498 @p_498 @p_498 @p_498) :rule resolution :premises (t527.t1 t527.t2 t527.t3 t527.t4 t527.t5 t527.t6))
(step t527.t8 (cl @p_498 @p_498 @p_498 @p_498 @p_498 @p_486) :rule reordering :premises (t527.t7))
(step t527.t9 (cl @p_498 @p_486) :rule contraction :premises (t527.t8))
(step t527.t10 (cl @p_499 @p_486) :rule resolution :premises (t527.t0 t527.t9))
(step t527.t11 (cl @p_499 @p_490) :rule implies_neg2)
(step t527.t12 (cl @p_499 @p_499) :rule resolution :premises (t527.t10 t527.t11))
(step t527.t13 (cl @p_499) :rule contraction :premises (t527.t12))
(step t527.t14 (cl @p_498 @p_486) :rule implies :premises (t527.t13))
(step t527.t15 (cl @p_495 @p_122 @p_312 @p_437 @p_487 @p_435) :rule and_neg)
(step t527.t16 (cl @p_495) :rule resolution :premises (t527.t15 t298 t527.a2 t527.a0 t527.a4 t527.a3))
(step t527.t17 (cl @p_486) :rule resolution :premises (t527.t14 t527.t16))
(step t527 (cl @p_437 @p_122 @p_312 @p_435 @p_487 @p_486) :rule subproof :discharge (t527.a0 t527.a1 t527.a2 t527.a3 t527.a4))
(step t528 (cl (! (not @p_492) :named @p_493) @p_7) :rule and_pos :args (0))
(step t529 (cl @p_493 @p_19) :rule and_pos :args (1))
(step t530 (cl @p_493 @p_79) :rule and_pos :args (2))
(step t531 (cl @p_493 @p_434) :rule and_pos :args (3))
(step t532 (cl @p_493 @p_438) :rule and_pos :args (4))
(step t533 (cl @p_486 @p_493 @p_493 @p_493 @p_493 @p_493) :rule resolution :premises (t527 t528 t529 t530 t531 t532))
(step t534 (cl @p_493 @p_493 @p_493 @p_493 @p_493 @p_486) :rule reordering :premises (t533))
(step t535 (cl @p_493 @p_486) :rule contraction :premises (t534))
(step t536 (cl @p_494 @p_486) :rule resolution :premises (t526 t535))
(step t537 (cl @p_494 @p_490) :rule implies_neg2)
(step t538 (cl @p_494 @p_494) :rule resolution :premises (t536 t537))
(step t539 (cl @p_494) :rule contraction :premises (t538))
(step t540 (cl @p_493 @p_486) :rule implies :premises (t539))
(step t541 (cl @p_437 @p_122 @p_312 @p_435 @p_487 @p_486) :rule resolution :premises (t525 t540))
(step t542 (cl @p_489 (not @p_437)) :rule or_neg :args (0))
(step t543 (cl @p_489 @p_125) :rule or_neg :args (1))
(step t544 (cl @p_489 @p_318) :rule or_neg :args (2))
(step t545 (cl @p_489 @p_491) :rule or_neg :args (3))
(step t546 (cl @p_489 (not @p_487)) :rule or_neg :args (4))
(step t547 (cl @p_489 @p_490) :rule or_neg :args (5))
(step t548 (cl @p_489 @p_489 @p_489 @p_489 @p_489 @p_489) :rule resolution :premises (t541 t542 t543 t544 t545 t546 t547))
(step t549 (cl @p_489) :rule contraction :premises (t548))
(step t550 (cl @p_488) :rule resolution :premises (t521 t524 t549))
(step t551 (cl @p_437 @p_18 @p_312 @p_435 @p_487 @p_486) :rule or :premises (t550))
(step t552 (cl @p_18 @p_312 @p_486 @p_435 @p_437 @p_487) :rule reordering :premises (t551))
(step t553 (cl (! (=> (! (and @p_13 @p_7) :named @p_468) false) :named @p_470) @p_468) :rule implies_neg1)
(anchor :step t554)
(assume t554.a0 @p_13)
(assume t554.a1 @p_7)
(step t554.t0 (cl (not (! (= (! (<= (! (+ @p_6 @p_140) :named @p_455) (! (+ 0 (! (* -1 2) :named @p_471)) :named @p_472)) :named @p_473) false) :named @p_476)) (not @p_473) false) :rule equiv_pos2)
(step t554.t1 (cl (! (= @p_455 0) :named @p_467)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_467 3 7))
(step t554.t2 (cl (! (= @p_471 -2) :named @p_479)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_479 3 7))
(step t554.t3 (cl (= @p_472 (! (+ 0 -2) :named @p_477))) :rule cong :premises (t129 t554.t2))
(step t554.t4 (cl (! (= @p_477 -2) :named @p_478)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_478 3 7))
(step t554.t5 (cl (= @p_472 -2)) :rule trans :premises (t554.t3 t554.t4))
(step t554.t6 (cl (= @p_473 @p_284)) :rule cong :premises (t554.t1 t554.t5))
(step t554.t7 (cl @p_285) :rule hole :args ("TRUST_THEORY_REWRITE" @p_285 3 7))
(step t554.t8 (cl @p_476) :rule trans :premises (t554.t6 t554.t7))
(step t554.t9 (cl @p_437 (not (! (= @p_140 @p_471) :named @p_474)) @p_473) :rule la_generic :args (-1/1 -1/1 1/1))
(step t554.t10 (cl (=> (! (and @p_147 @p_13) :named @p_475) @p_474)) :rule la_mult_neg)
(step t554.t11 (cl (not @p_475) @p_474) :rule implies :premises (t554.t10))
(step t554.t12 (cl @p_475 @p_150 @p_253) :rule and_neg)
(step t554.t13 (cl (= @p_149 @p_147)) :rule equiv_simplify)
(step t554.t14 (cl (not @p_149) @p_147) :rule equiv1 :premises (t554.t13))
(step t554.t15 (cl @p_149) :rule rare_rewrite :args ("evaluate"))
(step t554.t16 (cl @p_147) :rule resolution :premises (t554.t14 t554.t15))
(step t554.t17 (cl @p_475) :rule resolution :premises (t554.t12 t554.t16 t554.a0))
(step t554.t18 (cl @p_474) :rule resolution :premises (t554.t11 t554.t17))
(step t554.t19 (cl @p_473) :rule resolution :premises (t554.t9 t554.a1 t554.t18))
(step t554.t20 (cl false) :rule resolution :premises (t554.t0 t554.t8 t554.t19))
(step t554 (cl @p_253 @p_437 false) :rule subproof :discharge (t554.a0 t554.a1))
(step t555 (cl (! (not @p_468) :named @p_469) @p_13) :rule and_pos :args (0))
(step t556 (cl @p_469 @p_7) :rule and_pos :args (1))
(step t557 (cl false @p_469 @p_469) :rule resolution :premises (t554 t555 t556))
(step t558 (cl @p_469 @p_469 false) :rule reordering :premises (t557))
(step t559 (cl @p_469 false) :rule contraction :premises (t558))
(step t560 (cl @p_470 false) :rule resolution :premises (t553 t559))
(step t561 (cl @p_470 @p_132) :rule implies_neg2)
(step t562 (cl @p_470 @p_470) :rule resolution :premises (t560 t561))
(step t563 (cl @p_470) :rule contraction :premises (t562))
(step t564 (cl (= @p_470 @p_469)) :rule implies_simplify)
(step t565 (cl (not @p_470) @p_469) :rule equiv1 :premises (t564))
(step t566 (cl @p_469) :rule resolution :premises (t563 t565))
(step t567 (cl @p_253 @p_437) :rule not_and :premises (t566))
(step t568 (cl (! (=> (! (and @p_10 @p_7) :named @p_450) false) :named @p_452) @p_450) :rule implies_neg1)
(anchor :step t569)
(assume t569.a0 @p_10)
(assume t569.a1 @p_7)
(step t569.t0 (cl (not (! (= (! (<= @p_455 (! (+ 0 (! (* -1 1) :named @p_453)) :named @p_454)) :named @p_456) false) :named @p_461)) (not @p_456) false) :rule equiv_pos2)
(step t569.t1 (cl @p_467) :rule hole :args ("TRUST_THEORY_REWRITE" @p_467 3 7))
(step t569.t2 (cl (! (= @p_453 -1) :named @p_466)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_466 3 7))
(step t569.t3 (cl (= @p_454 (! (+ 0 -1) :named @p_464))) :rule cong :premises (t129 t569.t2))
(step t569.t4 (cl (! (= @p_464 -1) :named @p_465)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_465 3 7))
(step t569.t5 (cl (= @p_454 -1)) :rule trans :premises (t569.t3 t569.t4))
(step t569.t6 (cl (= @p_456 (! (<= 0 -1) :named @p_462))) :rule cong :premises (t569.t1 t569.t5))
(step t569.t7 (cl (! (= @p_462 false) :named @p_463)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_463 3 7))
(step t569.t8 (cl @p_461) :rule trans :premises (t569.t6 t569.t7))
(step t569.t9 (cl @p_437 (not (! (= @p_140 @p_453) :named @p_457)) @p_456) :rule la_generic :args (-1/1 -1/1 1/1))
(step t569.t10 (cl (=> (! (and @p_147 @p_10) :named @p_458) @p_457)) :rule la_mult_neg)
(step t569.t11 (cl (not @p_458) @p_457) :rule implies :premises (t569.t10))
(step t569.t12 (cl @p_458 @p_150 @p_449) :rule and_neg)
(step t569.t13 (cl (= @p_149 @p_147)) :rule equiv_simplify)
(step t569.t14 (cl (not @p_149) @p_147) :rule equiv1 :premises (t569.t13))
(step t569.t15 (cl @p_149) :rule rare_rewrite :args ("evaluate"))
(step t569.t16 (cl @p_147) :rule resolution :premises (t569.t14 t569.t15))
(step t569.t17 (cl @p_458) :rule resolution :premises (t569.t12 t569.t16 t569.a0))
(step t569.t18 (cl @p_457) :rule resolution :premises (t569.t11 t569.t17))
(step t569.t19 (cl @p_456) :rule resolution :premises (t569.t9 t569.a1 t569.t18))
(step t569.t20 (cl false) :rule resolution :premises (t569.t0 t569.t8 t569.t19))
(step t569 (cl @p_449 @p_437 false) :rule subproof :discharge (t569.a0 t569.a1))
(step t570 (cl (! (not @p_450) :named @p_451) @p_10) :rule and_pos :args (0))
(step t571 (cl @p_451 @p_7) :rule and_pos :args (1))
(step t572 (cl false @p_451 @p_451) :rule resolution :premises (t569 t570 t571))
(step t573 (cl @p_451 @p_451 false) :rule reordering :premises (t572))
(step t574 (cl @p_451 false) :rule contraction :premises (t573))
(step t575 (cl @p_452 false) :rule resolution :premises (t568 t574))
(step t576 (cl @p_452 @p_132) :rule implies_neg2)
(step t577 (cl @p_452 @p_452) :rule resolution :premises (t575 t576))
(step t578 (cl @p_452) :rule contraction :premises (t577))
(step t579 (cl (= @p_452 @p_451)) :rule implies_simplify)
(step t580 (cl (not @p_452) @p_451) :rule equiv1 :premises (t579))
(step t581 (cl @p_451) :rule resolution :premises (t578 t580))
(step t582 (cl @p_449 @p_437) :rule not_and :premises (t581))
(step t583 (cl @p_437 @p_438) :rule ite2 :premises (t519))
(step t584 (cl @p_438 @p_437) :rule reordering :premises (t583))
(step t585 (cl @p_437) :rule resolution :premises (t552 t302 t298 t327 t478 t567 t582 t584))
(step t586 (cl @p_436) :rule resolution :premises (t520 t585))
(step t587 (cl @p_435) :rule resolution :premises (t508 t298 t302 t454 t586))
(step t588 (cl @p_10) :rule resolution :premises (t478 t587))
(step t589 (cl @p_449 @p_433) :rule ite2 :premises (t477))
(step t590 (cl @p_433 @p_449) :rule reordering :premises (t589))
(step t591 (cl @p_433) :rule resolution :premises (t590 t588))
(step t592 (cl @p_532) :rule resolution :premises (t92 t111 t302 t454 t588 t591))
(step t593 (cl (! (=> @p_10 @p_515) :named @p_516) @p_10) :rule implies_neg1)
(anchor :step t594)
(assume t594.a0 @p_10)
(step t594.t0 (cl @p_516 @p_10) :rule implies_neg1)
(anchor :step t594.t1)
(assume t594.t1.a0 @p_10)
(step t594.t1.t0 (cl (not @p_529) @p_531 @p_515) :rule equiv_pos2)
(step t594.t1.t1 (cl @p_529) :rule hole :args ("TRUST_THEORY_REWRITE" @p_529 3 6))
(step t594.t1.t2 (cl @p_530 @p_517 @p_518) :rule equiv_pos2)
(step t594.t1.t3 (cl @p_528) :rule symm :premises (t594.t1.t1))
(step t594.t1.t4 (cl (! (=> @p_514 false) :named @p_519) @p_514) :rule implies_neg1)
(anchor :step t594.t1.t5)
(assume t594.t1.t5.a0 @p_514)
(step t594.t1.t5.t0 (cl (not (! (= (! (<= @p_278 (! (+ @p_471 1) :named @p_520)) :named @p_521) false) :named @p_525)) (not @p_521) false) :rule equiv_pos2)
(step t594.t1.t5.t1 (cl @p_289) :rule hole :args ("TRUST_THEORY_REWRITE" @p_289 3 7))
(step t594.t1.t5.t2 (cl @p_479) :rule hole :args ("TRUST_THEORY_REWRITE" @p_479 3 7))
(step t594.t1.t5.t3 (cl (= @p_520 (! (+ -2 1) :named @p_526))) :rule cong :premises (t594.t1.t5.t2 t126))
(step t594.t1.t5.t4 (cl (! (= @p_526 -1) :named @p_527)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_527 3 7))
(step t594.t1.t5.t5 (cl (= @p_520 -1)) :rule trans :premises (t594.t1.t5.t3 t594.t1.t5.t4))
(step t594.t1.t5.t6 (cl (= @p_521 @p_462)) :rule cong :premises (t594.t1.t5.t1 t594.t1.t5.t5))
(step t594.t1.t5.t7 (cl @p_463) :rule hole :args ("TRUST_THEORY_REWRITE" @p_463 3 7))
(step t594.t1.t5.t8 (cl @p_525) :rule trans :premises (t594.t1.t5.t6 t594.t1.t5.t7))
(step t594.t1.t5.t9 (cl (not (! (<= @p_140 @p_471) :named @p_522)) @p_449 @p_521) :rule la_generic :args (1/1 -1/1 1/1))
(step t594.t1.t5.t10 (cl (=> (! (and @p_147 @p_514) :named @p_523) @p_522)) :rule la_mult_neg)
(step t594.t1.t5.t11 (cl (not @p_523) @p_522) :rule implies :premises (t594.t1.t5.t10))
(step t594.t1.t5.t12 (cl @p_523 @p_150 @p_515) :rule and_neg)
(step t594.t1.t5.t13 (cl (= @p_149 @p_147)) :rule equiv_simplify)
(step t594.t1.t5.t14 (cl (not @p_149) @p_147) :rule equiv1 :premises (t594.t1.t5.t13))
(step t594.t1.t5.t15 (cl @p_149) :rule rare_rewrite :args ("evaluate"))
(step t594.t1.t5.t16 (cl @p_147) :rule resolution :premises (t594.t1.t5.t14 t594.t1.t5.t15))
(step t594.t1.t5.t17 (cl @p_523) :rule resolution :premises (t594.t1.t5.t12 t594.t1.t5.t16 t594.t1.t5.a0))
(step t594.t1.t5.t18 (cl @p_522) :rule resolution :premises (t594.t1.t5.t11 t594.t1.t5.t17))
(step t594.t1.t5.t19 (cl @p_521) :rule resolution :premises (t594.t1.t5.t9 t594.t1.t5.t18 t594.t1.a0))
(step t594.t1.t5.t20 (cl false) :rule resolution :premises (t594.t1.t5.t0 t594.t1.t5.t8 t594.t1.t5.t19))
(step t594.t1.t5 (cl @p_515 false) :rule subproof :discharge (t594.t1.t5.a0))
(step t594.t1.t6 (cl @p_519 false) :rule resolution :premises (t594.t1.t4 t594.t1.t5))
(step t594.t1.t7 (cl @p_519 @p_132) :rule implies_neg2)
(step t594.t1.t8 (cl @p_519 @p_519) :rule resolution :premises (t594.t1.t6 t594.t1.t7))
(step t594.t1.t9 (cl @p_519) :rule contraction :premises (t594.t1.t8))
(step t594.t1.t10 (cl (= @p_519 @p_515)) :rule implies_simplify)
(step t594.t1.t11 (cl (not @p_519) @p_515) :rule equiv1 :premises (t594.t1.t10))
(step t594.t1.t12 (cl @p_515) :rule resolution :premises (t594.t1.t9 t594.t1.t11))
(step t594.t1.t13 (cl @p_518) :rule resolution :premises (t594.t1.t2 t594.t1.t3 t594.t1.t12))
(step t594.t1.t14 (cl @p_515) :rule resolution :premises (t594.t1.t0 t594.t1.t1 t594.t1.t13))
(step t594.t1 (cl @p_449 @p_515) :rule subproof :discharge (t594.t1.a0))
(step t594.t2 (cl @p_516 @p_515) :rule resolution :premises (t594.t0 t594.t1))
(step t594.t3 (cl @p_516 @p_517) :rule implies_neg2)
(step t594.t4 (cl @p_516 @p_516) :rule resolution :premises (t594.t2 t594.t3))
(step t594.t5 (cl @p_516) :rule contraction :premises (t594.t4))
(step t594.t6 (cl @p_449 @p_515) :rule implies :premises (t594.t5))
(step t594.t7 (cl @p_515) :rule resolution :premises (t594.t6 t594.a0))
(step t594 (cl @p_449 @p_515) :rule subproof :discharge (t594.a0))
(step t595 (cl @p_516 @p_515) :rule resolution :premises (t593 t594))
(step t596 (cl @p_516 @p_517) :rule implies_neg2)
(step t597 (cl @p_516 @p_516) :rule resolution :premises (t595 t596))
(step t598 (cl @p_516) :rule contraction :premises (t597))
(step t599 (cl @p_449 @p_515) :rule implies :premises (t598))
(step t600 (cl @p_515) :rule resolution :premises (t599 t588))
(step t601 (cl @p_57) :rule and :premises (t110) :args (0))
(step t602 (cl) :rule resolution :premises (t57 t592 t600 t591 t454 t302 t601 t298))
